High-altitude hypoxia exposure inhibits erythrophagocytosis by inducing macrophage ferroptosis in the spleen

Author Response

The following is the authors’ response to the previous reviews.

Recommendations for the authors:

(1) Substantial revision of the claims and interpretation of the results is needed, especially in the setting of additional data showing enhanced erythrophagocytosis with decreased RBC lifespan.

Thank you for your valuable feedback and suggestion for a substantial revision of the claims and interpretation of our results. We acknowledge the importance of considering additional data that shows enhanced erythrophagocytosis with decreased RBC lifespan. In response, we have revised our manuscript and incorporated additional experimental data to support and clarify our findings.

(1) In our original manuscript, we reported a decrease in the number of splenic red pulp macrophages (RPMs) and phagocytic erythrocytes after hypobaric hypoxia (HH) exposure. This conclusion was primarily based on our observations of reduced phagocytosis in the spleen.

(2) Additional experimental data on RBC labeling and erythrophagocytosis:

Experiment 1 (RBC labeling and HH exposure)

We conducted an experiment where RBCs from mice were labeled with PKH67 and injected back into the mice. These mice were then exposed to normal normoxia (NN) or HH for 7 or 14 days. The subsequent assessment of RPMs in the spleen using flow cytometry and immunofluorescence detection revealed a significant decrease in both the population of splenic RPMs (F4/80hiCD11blo, new Figure 5A and C) and PKH67-positive macrophages after HH exposure (as depicted in new Figure 5A and C-E). This finding supports our original claim of reduced phagocytosis under HH conditions.

-Experiment 2 (erythrophagocytosis enhancement)

To examine the effects of enhanced erythrophagocytosis, we injected Tuftsin after administering PKH67-labelled RBCs. Our observations showed a significant decrease in PKH67 fluorescence in the spleen, particularly after Tuftsin injection compared to the NN group. This result suggests a reduction in RBC lifespan when erythrophagocytosis is enhanced (illustrated in new Figure 7, A-B).

(3) Revised conclusions:

The additional data from these experiments support our original findings by providing a more comprehensive view of the impact of HH exposure on splenic erythrophagocytosis.

The decrease in phagocytic RPMs and phagocytic erythrocytes after HH exposure, along with the observed decrease in RBC lifespan following enhanced erythrophagocytosis, collectively suggest a more complex interplay between hypoxia, erythrophagocytosis, and RBC lifespan than initially interpreted.

We think that these revisions and additional experimental data provide a more robust and detailed understanding of the effects of HH on splenic erythrophagocytosis and RBCs lifespan. We hope that these changes adequately address the concerns raised and strengthen the conclusions drawn in our manuscript.

(2) F4/80 high; CD11b low are true RPMs which the cells which the authors are presenting, i.e. splenic monocytes / pre-RPMs. To discuss RPM function requires the presentation of these cells specifically rather than general cells in the proper area of the spleen.

Thank you for your feedback requesting a substantial revision of our claims and interpretation, particularly considering additional data showing enhanced erythrophagocytosis with decreased RBC lifespan. In response, we have thoroughly revised our manuscript and included new experimental data that further elucidate the effects of HH on RPMs and erythrophagocytosis.

(1) Re-evaluation of RPMs population after HH exposure:

Flow cytometry analysis (new Figure 3G, Figure 5A and B): We revisited the analysis of RPMs (F4/80hiCD11blo) in the spleen after 7 and 14 days of HH exposure. Our revised flow cytometry data consistently showed a significant decrease in the RPMs population post-HH exposure, reinforcing our initial findings.

In situ expression of RPMs (Figure S1, A-D):

We further confirmed the decreased population of RPMs through in situ co-staining with F4/80 and CD11b, and F4/80 and CD68, in spleen tissues. These results clearly demonstrated a significant reduction in F4/80hiCD11blo (Figure S1, A and B) and F4/80hiCD68hi (Figure S1, C and D) cells following HH exposure.

(2) Single-cell sequencing analysis of splenic RPMs:

We conducted a single-cell sequencing analysis of spleen samples post 7 days of HH exposure (Figure S2, A-C). This analysis revealed a notable shift in the distribution of RPMs, predominantly associated with Cluster 0 under NN conditions, to a reduced presence in this cluster after HH exposure.

Pseudo-time series analysis indicated a transition pattern change in spleen RPMs, with a shift from Cluster 2 and Cluster 1 towards Cluster 0 under NN conditions, and a reverse transition following HH exposure (Figure S2, B and D). This finding implies a decrease in resident RPMs in the spleen under HH conditions.

(3) Consolidated findings and revised interpretation:

The comprehensive analysis of flow cytometry, in situ staining, and single-cell sequencing data consistently indicates a significant reduction in the number of RPMs following HH exposure.

These findings, taken together, strongly support the revised conclusion that HH exposure leads to a decrease in RPMs in the spleen, which in turn may affect erythrophagocytosis and RBC lifespan.

In conclusion, our revised manuscript now includes additional experimental data and analyses, strengthening our claims and providing a more nuanced interpretation of the impact of HH on spleen RPMs and related erythrophagocytosis processes. We believe these revisions and additional data address your concerns and enhance the scientific validity of our study.

(3) RBC retention in the spleen should be measured anyway quantitatively, eg, with proper flow cytometry, to determine whether it is increased or decreased.

Thank you for your query regarding the quantitative measurement of RBC retention in the spleen, particularly in relation to HH exposure. We have utilized a combination of techniques, including flow cytometry and histological staining, to investigate this aspect comprehensively. Below is a summary of our findings and methodology.

(1) Flow cytometry analysis of labeled RBCs:

Our study employed both NHS-biotin (new Figure 4, A-D) and PKH67 labeling (new Figure 4, E-H) to track RBCs in mice exposed to HH. Flow cytometry results from these experiments (new Figure 4, A-H) showed a decrease in the proportion of labeled RBCs over time, both in the blood and spleen. Notably, there was a significantly greater reduction in the amplitude of fluorescently labeled RBCs after NN exposure compared to the reduced amplitude of fluorescently labeled RBCs observed in blood and spleen under HH exposure. The observed decrease in labeled RBCs was initially counterintuitive, as we expected an increase in RBC retention due to reduced erythrophagocytosis. However, this decrease can be attributed to the significantly increased production of RBCs following HH exposure, diluting the proportion of labeled cells.

Specifically, for blood, the biotin-labeled RBCs decreased by 12.06% under NN exposure and by 7.82% under HH exposure, while the PKH67-labeled RBCs decreased by 9.70% under NN exposure and by 4.09% under HH exposure. For spleen, the biotin-labeled RBCs decreased by 3.13% under NN exposure and by 0.46% under HH exposure, while the PKH67-labeled RBCs decreased by 1.16% under NN exposure and by 0.92% under HH exposure. These findings suggest that HH exposure leads to a decrease in the clearance rate of RBCs.

(2) Detection of erythrophagocytosis in spleen:

To assess erythrophagocytosis directly, we labeled RBCs with PKH67 and analyzed their uptake by splenic macrophages (F4/80hi) after HH exposure. Our findings (new Figure 5, D-E) indicated a decrease in PKH67-positive macrophages in the spleen, suggesting reduced erythrophagocytosis.

(3) Flow cytometry analysis of RBC retention:

Our flow cytometry analysis revealed a decrease in PKH67-positive RBCs in both blood and spleen (Figure S4). We postulated that this was due to increased RBC production after HH exposure. However, this method might not accurately reflect RBC retention, as it measures the proportion of PKH67-labeled RBCs relative to the total number of RBCs, which increased after HH exposure.

(4) Histological and immunostaining analysis:

Histological examination using HE staining and band3 immunostaining in situ (new Figure 6, A-D, and G-H) revealed a significant increase in RBC numbers in the spleen after HH exposure. This was further confirmed by detecting retained RBCs in splenic single cells using Wright-Giemsa composite stain (new Figure 6, E and F) and retained PKH67-labelled RBCs in spleen (new Figure 6, I and J).

(5) Interpreting the data:

The comprehensive analysis suggests a complex interplay between increased RBC production and decreased erythrophagocytosis in the spleen following HH exposure. While flow cytometry indicated a decrease in the proportion of labeled RBCs, histological and immunostaining analyses demonstrated an actual increase in RBCs retention in the spleen. These findings collectively suggest that while the overall RBCs production is upregulated following HH exposure, the spleen's capacity for erythrophagocytosis is concurrently diminished, leading to increased RBCs retention.

(6) Conclusion:

Taken together, our results indicate a significant increase in RBCs retention in the spleen post-HH exposure, likely due to reduced residual RPMs and erythrophagocytosis. This conclusion is supported by a combination of flow cytometry, histological staining, and immunostaining techniques, providing a comprehensive view of RBC dynamics under HH conditions. We think these findings offer a clear quantitative measure of RBC retention in the spleen, addressing the concerns raised in your question.

(4) Numerous other methodological problems as listed below.

We appreciate your question, which highlights the importance of using multiple analytical approaches to understand complex physiological processes. Please find below our point-by-point response to the methodological comments.

Reviewer #1 (Recommendations For The Authors):

(1) Decreased BM and spleen monocytes d/t increased liver monocyte migration is unclear. there is no evidence that this happens or why it would be a reasonable hypothesis, even in splenectomized mice.

Thank you for highlighting the need for further clarification and justification of our hypothesized decrease in BM and spleen monocytes due to increased monocyte migration to the liver, particularly in the context of splenectomized mice. Indeed, our study has not explicitly verified an augmentation in mononuclear cell migration to the liver in splenectomized mice.

Nonetheless, our investigations have revealed a notable increase in monocyte migration to the liver after HH exposure. Noteworthy is our discovery of a significant upregulation in colony stimulating factor-1 (CSF-1) expression in the liver, observed after both 7 and 14 days of HH exposure (data not included). This observation was substantiated through flow cytometry analysis (as depicted in Figure S4), which affirmed an enhanced migration of monocytes to the liver. Specifically, we noted a considerable increase in the population of transient macrophages, monocytes, and Kupffer cells in the liver following HH exposure.

Considering these findings, we hypothesize that hypoxic conditions may activate a compensatory mechanism that directs monocytes towards the liver, potentially linked to the liver’s integral role in the systemic immune response. In accordance with these insights, we intend to revise our manuscript to reflect the speculative nature of this hypothesis more accurately, and to delineate the strategies we propose for its further empirical investigation. This amendment ensures that our hypothesis is presented with full consideration of its speculative basis, supported by a coherent framework for future validation.

(2) While F4/80+CD11b+ population is decreased, this is mainly driven by CD11b and F4/80+ alone population is significantly increased. This is counter to the hypothesis.

Thank you for addressing the apparent discrepancy in our findings concerning the F4/80+CD11b+ population and the increase in the F4/80+ alone population, which seems to contradict our initial hypothesis. Your observation is indeed crucial for the integrity of our study, and we appreciate the opportunity to clarify this matter.

(1) Clarification of flow cytometry results:

In response to the concerns raised, we revisited our flow cytometry experiments with a focus on more clearly distinguishing the cell populations. Our initial graph had some ambiguities in cell grouping, which might have led to misinterpretations.

The revised flow cytometry analysis, specifically aimed at identifying red pulp macrophages (RPMs) characterized as F4/80hiCD11blo in the spleen, demonstrated a significant decrease in the F4/80 population. This finding is now in alignment with our immunofluorescence results.

(2) Revised data and interpretation:

The results presented in new Figure 3G and Figure 5 (A and B) consistently indicate a notable reduction in the RPMs population following HH exposure. This supports our revised understanding that HH exposure leads to a decrease in the specific macrophage subset (F4/80hiCD11blo) in the spleen.

We’ve updated our manuscript to reflect these new findings and interpretations. The revised manuscript details the revised flow cytometry analysis and discusses the potential mechanisms behind the observed changes in macrophage populations.

(3) HO-1 expression cannot be used as a surrogate to quantify number of macrophages as the expression per cell can decrease and give the same results. In addition, the localization of effect to the red pulp is not equivalent to an assertion that the conclusion applies to macrophages given the heterogeneity of this part of the organ and the spleen in general.

Thank you for your insightful comments regarding the use of HO-1 expression as a surrogate marker for quantifying macrophage numbers, and for pointing out the complexity of attributing changes in HO-1 expression specifically to macrophages in the splenic red pulp. Your observations are indeed valid and warrant a detailed response.

(1) Role of HO-1 in macrophage activity:

In our study, HO-1 expression was not utilized as a direct marker for quantifying macrophages. Instead, it was considered an indicator of macrophage activity, particularly in relation to erythrophagocytosis. HO-1, being upregulated in response to erythrophagocytosis, serves as an indirect marker of this process within splenic macrophages.

The rationale behind this approach was that increased HO-1 expression, induced by erythrophagocytosis in the spleen’s red pulp, could suggest an augmentation in the activity of splenic macrophages involved in this process.

(2) Limitations of using HO-1 as an indicator:

We acknowledge your point that HO-1 expression per cell might decrease, potentially leading to misleading interpretations if used as a direct quantifier of macrophage numbers. The variability in HO-1 expression per cell indeed presents a limitation in using it as a sole indicator of macrophage quantity.

Furthermore, your observation about the heterogeneity of the spleen, particularly the red pulp, is crucial. The red pulp is a complex environment with various cell types, and asserting that changes in HO-1 expression are exclusive to macrophages could oversimplify this complexity.

(3) Addressing the concerns:

To address these concerns, we propose to supplement our HO-1 expression data with additional specific markers for macrophages. This would help in correlating HO-1 expression more accurately with macrophage numbers and activity.

We also plan to conduct further studies to delineate the specific cell types in the red pulp contributing to HO-1 expression. This could involve techniques such as immunofluorescence or immunohistochemistry, which would allow us to localize HO-1 expression to specific cell populations within the splenic red pulp.

We’ve revised our manuscript to clarify the role of HO-1 expression as an indirect marker of erythrophagocytosis and to acknowledge its limitations as a surrogate for quantifying macrophage numbers.

(4) line 63-65 is inaccurate as red cell homeostasis reaches a new steady state in chronic hypoxia.

Thank you for pointing out the inaccuracy in lines 63-65 of our manuscript regarding red cell homeostasis in chronic hypoxia. Your feedback is invaluable in ensuring the accuracy and scientific integrity of our work. We’ve revised lines 63-65 to accurately reflect the understanding.

(5) Eryptosis is not defined in the manuscript.

Thank you for highlighting the omission of a definition for eryptosis in our manuscript. We acknowledge the significance of precisely defining such key terminologies, particularly when they play a crucial role in the context of our research findings. Eryptosis, a term referenced in our study, is a specialized form of programmed cell death unique to erythrocytes. Similar with apoptosis in other cell types, eryptosis is characterized by distinct physiological changes including cell shrinkage, membrane blebbing, and the externalization of phosphatidylserine on the erythrocyte surface. These features are indicative of the RBCs lifecycle and its regulated destruction process.

However, it is pertinent to note that our current study does not extensively delve into the mechanisms or implications of eryptosis. Our primary focus has been to elucidate the effects of HH exposure on the processes of splenic erythrophagocytosis and the resultant impact on the lifespan of RBCs. Given this focus, and to maintain the coherence and relevance of our manuscript, we have decided to exclude specific discussions of eryptosis from our revised manuscript. This decision aligns with our aim to provide a clear and concentrated exploration of the influence of HH exposure on RBCs dynamics and splenic function.

We appreciate your input, which has significantly contributed to enhancing the clarity and accuracy of our manuscript. The revision ensures that our research is presented with a focused scope, aligning closely with our experimental investigations and findings.

(6) Physiologically, there is no evidence that there is any "free iron" in cells, making line 89 point inaccurate.

Thank you for highlighting the concern regarding the reference to "free iron" in cells in line 89 of our manuscript. The term "free iron" in our manuscript was intended to refer to divalent iron (Fe2+), rather than unbound iron ions freely circulating within cells. We acknowledge that the term "free iron" might lead to misconceptions, as it implies the presence of unchelated iron, which is not physiologically common due to the potential for oxidative damage. To rectify this and provide clarity, we’ve revised line 89 of our manuscript to reflect our meaning more accurately. Instead of "free iron," we use "divalent iron (Fe2+)" to avoid any misunderstanding regarding the state of iron in cells. We also ensure that any implications drawn from the presence of Fe2+ in cells are consistent with current scientific literature and understanding.

(7) Fig 1f no stats

We appreciate your critical review and suggestions, which help in improving the accuracy and clarity of our research. We’ve revised statistic diagram of new Figure 1F.

(8) Splenectomy experiments demonstrate that erythrophagocytosis is almost completely replaced by functional macrophages in other tissues (likely Kupffer cells in the liver). there is only a minor defect and no data on whether it is in fact the liver or other organs that provide this replacement function and makes the assertions in lines 345-349 significantly overstated.

Thank you for your critical assessment of our interpretation of the splenectomy experiments, especially concerning the role of erythrophagocytosis by macrophages in other tissues, such as Kupffer cells in the liver. We appreciate your observation that our assertions may be overstated and acknowledge the need for more specific data to identify which organs compensate for the loss of splenic erythrophagocytosis.

(1) Splenectomy experiment findings:

Our findings in Figure 2D do indicate that in the splenectomized group under NN conditions, erythrophagocytosis is substantially compensated for by functional macrophages in other tissues. This is an important observation that highlights the body's ability to adapt to the loss of splenic function.

However, under HH conditions, our data suggest that the spleen plays an important role in managing erythrocyte turnover, as indicated by the significant impact of splenectomy on erythrophagocytosis and subsequent erythrocyte dynamics.

(2) Addressing the lack of specific organ identification:

We acknowledge that our study does not definitively identify which organs, such as the liver or others, take over the erythrophagocytosis function post-splenectomy. This is an important aspect that needs further investigation.

To address this, we also plan to perform additional experiments that could more accurately point out the specific tissues compensating for the loss of splenic erythrophagocytosis. This could involve tracking labeled erythrocytes or using specific markers to identify macrophages actively engaged in erythrophagocytosis in various organs.

(3) Revising manuscript statements:

Considering your feedback, we’ve revised the statements in lines 345-349 (lines 378-383 in revised manuscript) to enhance the scientific rigor and clarity of our research presentation.

(9) M1 vs M2 macrophage experiments are irrelevant to the main thrust of the manuscript, there are no references to support the use of only CD16 and CD86 for these purposes, and no stats are provided. It is also unclear why bone marrow monocyte data is presented and how it is relevant to the rest of the manuscript.

Thank you for your critical evaluation of the relevance and presentation of the M1 vs. M2 macrophage experiments in our manuscript. We appreciate your insights, especially regarding the use of specific markers and the lack of statistical analysis, as well as the relevance of bone marrow monocyte data to our study's main focus.

(1) Removal of M1 and M2 macrophage data:

Based on your feedback and our reassessment, we agree that the results pertaining to M1 and M2 macrophages did not align well with the main objectives of our manuscript. Consequently, we have decided to remove the related content on M1 and M2 macrophages from the revised manuscript. This decision was made to ensure that our manuscript remains focused and coherent, highlighting our primary findings without the distraction of unrelated or insufficiently supported data.

The use of only CD16 and CD86 markers for M1 and M2 macrophage characterization, without appropriate statistical analysis, was indeed a methodological limitation. We recognize that a more comprehensive set of markers and rigorous statistical analysis would be necessary for a meaningful interpretation of M1/M2 macrophage polarization. Furthermore, the relevance of these experiments to the central theme of our manuscript was not adequately established. Our study primarily focuses on erythrophagocytosis and red pulp macrophage dynamics under hypobaric hypoxia, and the M1/M2 polarization aspect did not contribute significantly to this narrative.

(2) Clarification on bone marrow monocyte data:

Regarding the inclusion of bone marrow monocyte data, we acknowledge that its relevance to the main thrust of the manuscript was not clearly articulated. In the revised manuscript, we provide a clearer rationale for its inclusion and how it relates to our primary objectives.

(3) Commitment to clarity and relevance:

We are committed to ensuring that every component of our manuscript contributes meaningfully to our overall objectives and research questions. Your feedback has been instrumental in guiding us to streamline our focus and present our findings more effectively.

We appreciate your valuable feedback, which has led to a more focused and relevant presentation of our research. These changes enhance the clarity and impact of our manuscript, ensuring that it accurately reflects our key research findings.

(10) Biotinolated RBC clearance is enhanced, demonstrating that RBC erythrophagocytosis is in fact ENHANCED, not diminished, calling into question the founding hypothesis that the manuscript proposes.

Thank you for your critical evaluation of our data on biotinylated RBC clearance, which suggests enhanced erythrophagocytosis under HH conditions. This observation indeed challenges our founding hypothesis that erythrophagocytosis is diminished in this setting. Below is a summary of our findings and methodology.

(1) Interpretation of RBC labeling results:

Both the previous results of NHS-biotin labeled RBCs (new Figure 4, A-D) and the current results of PKH67-labeled RBCs (new Figure 4, E-H) demonstrated a decrease in the number of labeled RBCs with an increase in injection time. The production of RBCs, including bone marrow and spleen production, was significantly increased following HH exposure, resulting in a consistent decrease in the proportion of labeled RBCs via flow cytometry detection both in the blood and spleen of mice compared to the NN group. However, compared to the reduced amplitude of fluorescently labeled RBCs observed in blood and spleen under NN exposure, there was a significantly weaker reduction in the amplitude of fluorescently labeled RBCs after HH exposure. Specifically, for blood, the biotin-labeled RBCs decreased by 12.06% under NN exposure and by 7.82% under HH exposure, while the PKH67-labeled RBCs decreased by 9.70% under NN exposure and by 4.09% under HH exposure. For spleen, the biotin-labeled RBCs decreased by 3.13% under NN exposure and by 0.46% under HH exposure, while the PKH67-labeled RBCs decreased by 1.16% under NN exposure and by 0.92% under HH exposure.

(2) Increased RBCs production under HH conditions:

It's important to note that RBCs production, including from bone marrow and spleen, was significantly increased following HH exposure. This increase in RBCs production could contribute to the decreased proportion of labeled RBCs observed in flow cytometry analyses, as there are more unlabeled RBCs diluting the proportion of labeled cells in the blood and spleen.

(3) Analysis of erythrophagocytosis in RPMs:

Our analysis of PKH67-labeled RBCs content within RPMs following HH exposure showed a significant reduction in the number of PKH67-positive RPMs in the spleen (new Figure 5). This finding suggests a decrease in erythrophagocytosis by RPMs under HH conditions.

(4) Reconciling the findings:

The apparent contradiction between enhanced RBC clearance (suggested by the reduced proportion of labeled RBCs) and reduced erythrophagocytosis in RPMs (indicated by fewer PKH67-positive RPMs) may be explained by the increased overall production of RBCs under HH. This increased production could mask the actual erythrophagocytosis activity in terms of the proportion of labeled cells. Therefore, while the proportion of labeled RBCs decreases more significantly under HH conditions, this does not necessarily indicate an enhanced erythrophagocytosis rate, but rather an increased dilution effect due to higher RBCs turnover.

(5) Revised interpretation and manuscript changes:

Given these factors, we update our manuscript to reflect this detailed interpretation and clarify the implications of the increased RBCs production under HH conditions on our observations of labeled RBCs clearance and erythrophagocytosis. We appreciate your insightful feedback, which has prompted a careful re-examination of our data and interpretations. We hope that these revisions provide a more accurate and comprehensive understanding of the effects of HH on erythrophagocytosis and RBCs dynamics.

(11) Legend in Fig 4c-4d looks incorrect and Fig 4e-4f is very non-specific since Wright stain does not provide evidence of what type of cells these are and making for a significant overstatement in the contribution of this data to "confirming" increased erythrophagocytosis in the spleen under HH exposure (line 395-396).

Thank you for your insightful observations regarding the data presentation and figure legends in our manuscript, particularly in relation to Figure 4 (renamed as Figure 6 in the revised manuscript) and the use of Wright-Giemsa composite staining. We appreciate your constructive feedback and acknowledge the importance of presenting our data with utmost clarity and precision.

(1) Amendments to Figure legends:

We recognize the necessity of rectifying inaccuracies in the legends of the previously labeled Figure 4C and D. Corrections have been meticulously implemented to ensure the legends accurately contain the data presented. Additionally, we acknowledge the error concerning the description of Wright staining. The method employed in our study is Wright-Giemsa composite staining, which, unlike Wright staining that solely stains cytoplasm (RBC), is capable of staining both nuclei and cytoplasm.

(2) Addressing the specificity of Wright-Giemsa Composite staining:

Our approach involved quantifying RBC retention using Wright-Giemsa composite staining on single splenic cells post-perfusion at 7 and 14 days post HH exposure. We understand and appreciate your concerns regarding the nonspecific nature of Wright staining. Although Wright stain is a general hematologic stain and not explicitly specific for certain cell types, its application in our study aimed to provide preliminary insights. The spleen cells, devoid of nuclei and thus likely to be RBCs, were stained and observed post-perfusion, indicating RBC retention within the spleen.

(3) Incorporating additional methods for RBC identification:

To enhance the specificity of our findings, we integrated supplementary methods for RBC identification in the revised manuscript. We employed band3 immunostaining (in the new Figure 6, C-D and G-H) and PKH67 labeling (Figure 6, I-J) for a more targeted identification of RBCs. Band3, serving as a reliable marker for RBCs, augments the specificity of our immunostaining approach. Likewise, PKH67 labeling affords a direct and definitive means to assess RBC retention in the spleen following HH exposure.

(4) Revised interpretation and manuscript modifications:

Based on these enhanced methodologies, we have refined our interpretation of the data and accordingly updated the manuscript. The revised narrative underscores that our conclusions regarding reduced erythrophagocytosis and RBC retention under HH conditions are corroborated by not only Wright-Giemsa composite staining but also by band3 immunostaining and PKH67 labeling, each contributing distinctively to our comprehensive understanding.

We are committed to ensuring that our manuscript precisely reflects the contribution of each method to our findings and conclusions. Your thorough review has been invaluable in identifying and rectifying areas for improvement in our research report and interpretation.

(12) Ferroptosis data in Fig 5 is not specific to macrophages and Fer-1 data confirms the expected effect of Fer-1 but there is no data that supports that Fer-1 reverses the destruction of these cells or restores their function in hypoxia. Finally, these experiments were performed in peritoneal macrophages which are functionally distinct from splenic RPM.

Thank you for your critique of our presentation and interpretation of the ferroptosis data in Figure 5 (renamed as Figure 9 in the revised manuscript), as well as your observations regarding the specificity of the experiments to macrophages and the effects of Fer-1. We value your input and acknowledge the need to clarify these aspects in our manuscript.

(1) Clarification on cell type used in experiments:

We appreciate your attention to the details of our experimental setup. The experiments presented in Figure 9 were indeed conducted on splenic macrophages, not peritoneal macrophages, as incorrectly mentioned in the original figure legend. This was an error in our manuscript, and we have revised the figure legend accordingly to accurately reflect the cell type used.

(2) Specificity of ferroptosis data:

We recognize that the data presented in Figure 9 need to be more explicitly linked to the specific macrophage population being studied. In the revised manuscript, we ensure that the discussion around ferroptosis data is clearly situated within the framework of splenic macrophages.

We also provide additional methodological details in the 'Methods' section to reinforce the specificity of our experiments to splenic macrophages.

(3) Effects of Fer-1 on macrophage function and survival:

Regarding the effect of Fer-1, we agree that while our data confirms the expected effect of Fer-1 in inhibiting ferroptosis, we have not provided direct evidence that Fer-1 reverses the destruction of macrophages or restores their function in hypoxia.

To address this, we propose additional experiments to specifically investigate the impact of Fer-1 on the survival and functional restoration of splenic macrophages under hypoxic conditions. This would involve assessing not only the inhibition of ferroptosis but also the recovery of macrophage functionality post-treatment.

(4) Revised interpretation and manuscript changes:

We’ve revised the relevant sections of our manuscript to reflect these clarifications and proposed additional studies. This includes modifying the discussion of the ferroptosis data to more accurately represent the cell types involved and the limitations of our current findings regarding the effects of Fer-1.

The revised manuscript presents a more detailed interpretation of the ferroptosis data, clearly describing what our current experiments demonstrate and what remains to be investigated.

We are grateful for your insightful feedback, which has highlighted important areas for improvement in our research presentation. We think that these revisions will enhance the clarity and scientific accuracy of our manuscript, ensuring that our findings and conclusions are well-supported and precisely communicated.

Reviewer #2 (Recommendations For The Authors):

The following questions and remarks should be considered by the authors:

(1) The methods should clearly state whether the HH was discontinued during the 7 or 14 day exposure for cleaning, fresh water etc. Moreover, how was CO2 controlled? The procedure for splenectomy needs to be described in the methods.

Thank you for your inquiry regarding the specifics of our experimental methods, particularly the management of HH exposure and the procedure for splenectomy. We appreciate your attention to detail and the importance of these aspects for the reproducibility and clarity of our research.

(1) HH exposure conditions:

In our experiments, mice were continuously exposed to HH for the entire duration of 7 or 14 days, without interruption for activities such as cleaning or providing fresh water. This uninterrupted exposure was crucial for maintaining consistent hypobaric conditions throughout the experiment. The hypobaric chamber was configured to ensure a ventilation rate of 25 air exchanges per minute. This high ventilation rate was effective in regulating the concentration of CO2 inside the chamber, thereby maintaining a stable environment for the mice.

(2) The splenectomy was performed as follows:

After anesthesia, the mice were placed in a supine position, and their limbs were fixed. The abdominal operation area was skinned, disinfected, and covered with a sterile towel. A median incision was made in the upper abdomen, followed by laparotomy to locate the spleen. The spleen was then carefully pulled out through the incision. The arterial and venous directions in the splenic pedicle were examined, and two vascular forceps were used to clamp all the tissue in the main cadre of blood vessels below the splenic portal. The splenic pedicle was cut between the forceps to remove the spleen. The end of the proximal hepatic artery was clamped with a vascular clamp, and double or through ligation was performed to secure the site. The abdominal cavity was then cleaned to ensure there was no bleeding at the ligation site, and the incision was closed. Post-operatively, the animals were housed individually. Generally, they were able to feed themselves after recovering from anesthesia and did not require special care.

We hope this detailed description addresses your queries and provides a clear understanding of the experimental conditions and procedures used in our study. These methodological details are crucial for ensuring the accuracy and reproducibility of our research findings.

(2) The lack of changes in MCH needs explanation? During stress erythropoiesis some limit in iron availability should cause MCH decrease particularly if the authors claim that macrophages for rapid iron recycling are decreased. Fig 1A is dispensable. Fig 1G NN control 14 days does not make sense since it is higher than 7 days of HH.

Thank you for your inquiry regarding the lack of changes in Mean Corpuscular Hemoglobin (MCH) in our study, particularly in the context of stress erythropoiesis and decreased macrophage-mediated iron recycling. We appreciate the opportunity to provide further clarification on this aspect.

(1) Explanation for stable MCH levels:

Our research identified a decrease in erythrophagocytosis and iron recycling in the spleen following HH exposure. Despite this, the MCH levels remained stable. This observation can be explained by considering the compensatory roles of other organs, particularly the liver and duodenum, in maintaining iron homeostasis.

Specifically, our investigations revealed an enhanced capacity of the liver to engulf RBCs and process iron under HH conditions. This increased hepatic erythrophagocytosis likely compensates for the reduced splenic activity, thereby stabilizing MCH levels.

(2) Role of hepcidin and DMT1 expression:

Additionally, hypoxia is known to influence iron metabolism through the downregulation of Hepcidin and upregulation of Divalent Metal Transporter 1 (DMT1) expression. These alterations lead to enhanced intestinal iron absorption and increased blood iron levels, further contributing to the maintenance of MCH levels despite reduced splenic iron recycling.

(3) Revised Figure 1 and data presentation

To address the confusion regarding the data presented in Figure 1G, we have made revisions in our manuscript. The original Figure 1G, which did not align with the expected trends, has been removed. In its place, we have included a statistical chart of Figure 1F in the new version of Figure 1G. This revision will provide a clearer and more accurate representation of our findings.

(4) Manuscript updates and future research:

We update our manuscript to incorporate these explanations, ensuring that the rationale behind the stable MCH levels is clearly articulated. This includes a discussion on the role of the liver and duodenum in iron metabolism under hypoxic conditions.

Future research could explore in greater detail the mechanisms by which different organs contribute to iron homeostasis under stress conditions like HH, particularly focusing on the dynamic interplay between hepatic and splenic functions.

We thank you for your insightful question, which has prompted a thorough re-examination of our findings and interpretations. We believe that these clarifications will enhance the overall understanding of our study and its implications in the context of iron metabolism and erythropoiesis under hypoxic conditions.

(3) Fig 2 the difference between sham and splenectomy is really marginal and not convincing. Is there also a difference at 7 days? Why does the spleen size decrease between 7 and 14 days?

Thank you for your observations regarding the marginal differences observed between sham and splenectomy groups in Figure 2, as well as your inquiries about spleen size dynamics over time. We appreciate this opportunity to clarify these aspects of our study.

(1) Splenectomy vs. Sham group differences:

In our experiments, the difference between the sham and splenectomy groups under HH conditions, though subtle, was consistent with our hypothesis regarding the spleen's role in erythrophagocytosis and stress erythropoiesis. Under NN conditions, no significant difference was observed between these groups, which aligns with the expectation that the spleen's contribution is more pronounced under hypoxic stress.

(2) Spleen size dynamics and peak stress erythropoiesis:

The observed splenic enlargement prior to 7 days can be attributed to a combination of factors, including the retention of RBCs and extramedullary hematopoiesis, which is known to be a response to hypoxic stress.

Prior research has elucidated that splenic stress-induced erythropoiesis, triggered by hypoxic conditions, typically attains its zenith within a timeframe of 3 to 7 days. This observation aligns with our Toluidine Blue (TO) staining results, which indicated that the apex of this response occurs at the 7-day mark (as depicted in Figure 1, F-G). Here, the culmination of this peak is characteristically succeeded by a diminution in extramedullary hematopoiesis, a phenomenon that could elucidate the observed contraction in spleen size, particularly in the interval between 7 and 14 days.

This pattern of splenic response under prolonged hypoxic stress is corroborated by studies such as those conducted by Wang et al. (2021), Harada et al. (2015), and Cenariu et al. (2021). These references collectively underscore that the spleen undergoes significant dynamism in reaction to sustained hypoxia. This dynamism is initially manifested as an enlargement of the spleen, attributable to escalated erythropoiesis and erythrophagocytosis. Subsequently, as these processes approach normalization, a regression in spleen size ensues.

We’ve revised our manuscript to include a more detailed explanation of these splenic dynamics under HH conditions, referencing the relevant literature to provide a comprehensive context for our findings. We will also consider performing additional analysis or providing further data on spleen size changes at 7 days to support our observations and ensure a thorough understanding of the splenic response to hypoxic stress over time.

(4) Fig 3 B the clusters should be explained in detail. If the decrease in macrophages in Fig 3K/L is responsible for the effect, why does splenectomy not have a much stronger effect? How do the authors know which cells died in the calcein stained population in Fig 3D?

Thank you for your insightful questions regarding the details of our data presentation in Figure 3, particularly about the identification of cell clusters and the implications of macrophage reduction. We appreciate the opportunity to address these aspects and clarify our findings.

(1) Explanation of cell clusters in Figure 3B:

In the revised manuscript, we have included detailed notes for each cell population represented in Figure 3B (Figure 3D in revised manuscript). These notes provide a clearer understanding of the cell types present in each cluster, enhancing the interpretability of our single-cell sequencing data.

This detailed annotation will help readers to better understand the composition of the splenic cell populations under study and how they are affected by hypoxic conditions.

(2) Impact of splenectomy vs. macrophage reduction:

The interplay between the reduction in macrophage populations, as evidenced by our single-cell sequencing data, and the ramifications of splenectomy presents a multifaceted scenario. Notably, the observed decline in macrophage numbers following HH exposure does not straightforwardly equate to a comparable alteration in overall splenic function, as might be anticipated with splenectomy.

In the context of splenectomy under HH conditions, a significant escalation in the RBCs count was observed, surpassing that in non-splenectomized mice exposed to HH. This finding underscores the spleen's critical role in modulating RBCs dynamics under HH. It also indirectly suggests that the diminished phagocytic capacity of the spleen following HH exposure contributes to an augmented RBCs count, albeit to a lesser extent than in the splenectomy group. This difference is attributed to the fact that, while the number of RPMs in the spleen post-HH is reduced, they are still present, unlike in the case of splenectomy, where they are entirely absent.

Splenectomy entails the complete removal of the spleen, thus eliminating a broad spectrum of functions beyond erythrophagocytosis and iron recycling mediated by macrophages. The nuanced changes observed in our study may be reflective of the spleen's diverse functionalities and the organism's adaptive compensatory mechanisms in response to the loss of this organ.

(3) Calcein stained population in Figure 3D:

Regarding the identification of cell death in the calcein-stained population in Figure 3D (Figure 3A in revised manuscript), we acknowledge that the specific cell types undergoing death could not be distinctly determined from this analysis alone.

The calcein staining method allows for the visualization of live (calcein-positive) and dead (calcein-negative) cells, but it does not provide specific information about the cell types. The decrease in macrophage population was inferred from the single-cell sequencing data, which offered a more precise identification of cell types.

(4) Revised manuscript and data presentation:

Considering your feedback, we have revised our manuscript to provide a more comprehensive explanation of the data presented in Figure 3, including the nature of the cell clusters and the interpretation of the calcein staining results.

We have also updated the manuscript to reflect the removal of Figure 3K/L results and to provide a more focused discussion on the relevant findings.

We are grateful for your detailed review, which has helped us to refine our data presentation and interpretation. These clarifications and revisions will enhance the clarity and scientific rigor of our manuscript, ensuring that our conclusions are well-supported and accurately conveyed.

(5) Is the reduced phagocytic capacity in Fig 4B significant? Erythrophagocytosis is compromised due to the considerable spontaneous loss of labelled erythrocytes; could other assays help? (potentially by a modified Chromium release assay?). Is it necessary to stimulated phagocytosis to see a significant effect?

Thank you for your inquiry regarding the significance of the reduced phagocytic capacity observed in Figure 4B, and the potential for employing alternative assays to elucidate erythrophagocytosis dynamics under HH conditions.

(1) Significance of reduced phagocytic capacity:

The observed reduction in the amplitude of fluorescently labeled RBCs in both the blood and spleen under HH conditions suggests a decrease in erythrophagocytosis. This is indicative of a diminished phagocytic capacity, particularly when contrasted with NN conditions.

(2) Investigation of erythrophagocytosis dynamics:

To delve deeper into erythrophagocytosis under HH, we employed Tuftsin to enhance this process. Following the injection of PKH67-labeled RBCs and subsequent HH exposure, we noted a significant decrease in PKH67 fluorescence in the spleen, particularly marked after the administration of Tuftsin. This finding implies that stimulated erythrophagocytosis can influence RBCs lifespan.

(3) Erythrophagocytosis under normal and hypoxic conditions:

Under normal conditions, the reduction in phagocytic activity is less apparent without stimulation. However, under HH conditions, our findings demonstrate a clear weakening of the phagocytic effect. While we established that promoting phagocytosis under NN conditions affects RBC lifespan, the impact of enhanced phagocytosis under HH on RBCs numbers was not explicitly investigated.

(4) Potential for alternative assays:

Considering the considerable spontaneous loss of labeled erythrocytes, alternative assays such as a modified Chromium release assay could provide further insights. Such assays might offer a more nuanced understanding of erythrophagocytosis efficiency and the stability of labeled RBCs under different conditions.

(5) Future research directions:

The implications of these results suggest that future studies should focus on comparing the effects of stimulated phagocytosis under both NN and HH conditions. This would offer a clearer picture of the impact of hypoxia on the phagocytic capacity of macrophages and the subsequent effects on RBC turnover.

In summary, our findings indicate a diminished erythrophagocytic capacity, with enhanced phagocytosis affecting RBCs lifespan. Further investigation, potentially using alternative assays, would be beneficial to comprehensively understand the dynamics of erythrophagocytosis in different physiological states.

(6) Can the observed ferroptosis be influenced by bi- and not trivalent iron chelators?

Thank you for your question regarding the potential influence of bi- and trivalent iron chelators on ferroptosis under hypoxic conditions. We appreciate the opportunity to discuss the implications of our findings in this context.

(1) Analysis of iron chelators on ferroptosis:

In our study, we did not specifically analyze the effects of bi- and trivalent iron chelators on ferroptosis under hypoxia. However, our observations with Deferoxamine (DFO), a well-known iron chelator, provide some insights into how iron chelation may influence ferroptosis in splenic macrophages under hypoxic conditions.

(2) Effect of DFO on oxidative stress markers:

Our findings showed that under 1% O2, there was an increase in Malondialdehyde (MDA) content, a marker of lipid peroxidation, and a decrease in Glutathione (GSH) content, indicative of oxidative stress. These changes are consistent with the induction of ferroptosis, which is characterized by increased lipid peroxidation and depletion of antioxidants. Treatment with Ferrostatin-1 (Fer-1) and DFO effectively reversed these alterations. This suggests that DFO, like Fer-1, can mitigate ferroptosis in splenic macrophages under hypoxia, primarily by impacting MDA and GSH levels.

(3) Potential role of iron chelators in ferroptosis:

The effectiveness of DFO in reducing markers of ferroptosis indicates that iron availability plays a crucial role in the ferroptotic process under hypoxic conditions. It is plausible that both bi- and trivalent iron chelators could influence ferroptosis, given their ability to modulate iron availability within cells. Since ferroptosis is an iron-dependent form of cell death, chelating iron, irrespective of its valence state, could potentially disrupt the process by limiting the iron necessary for the generation of reactive oxygen species and lipid peroxidation.

(4) Additional research and manuscript updates:

Our study highlights the need for further research to explore the differential effects of various iron chelators on ferroptosis, particularly under hypoxic conditions. Such studies could provide a more comprehensive understanding of the role of iron in ferroptosis and the potential therapeutic applications of iron chelators. We update our manuscript to include these findings and discuss the potential implications of iron chelation in the context of ferroptosis under hypoxic conditions. This will provide a broader perspective on our research and its significance in understanding the mechanisms of ferroptosis.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: High-altitude polycythemia (HAPC) is a chronic condition prevalent in individuals residing at high altitudes (HA), characterized by an augmented production of red blood cells (RBCs) due to hypoxic conditions. Despite its prevalence, the pathophysiological basis and molecular mechanisms driving HAPC remain largely unelucidated. In our study, we utilized a mouse model exposed to hypobaric hypoxia (HH), replicating the environmental conditions experienced at 6000 m above sea level, coupled with in vitro analysis of primary splenic macrophages under 1% O2 to investigate these mechanisms. Our findings revealed that HH exposure significantly stimulates erythropoiesis, leading to erythrocytosis. This was accompanied by a notable sequence of splenic changes, initially characterized by splenic contraction, which subsequently progressed to splenomegaly over 14 days. A key observation was the impairment on the capacity of spleen to process RBCs, primarily attributed to a reduction in splenic macrophages located in the red pulp. Extended HH treatment over 7 and 14 days precipitated increased iron mobilization and the onset of ferroptosis within the spleen. This was evidenced by altered expression levels of iron metabolism and ferroptosis-related proteins, paralleling gene expression patterns observed in human peripheral blood mononuclear cells. Single-cell sequencing of splenic tissue post-HH exposure demonstrated a marked decrease in macrophage populations after 7 days. Our study further disclosed a significant increase in RBCs retention in the spleen following HH exposure, likely a consequence of diminished red pulp macrophages (RPMs) and erythrophagocytosis. This hypothesis was corroborated through comprehensive analyses involving flow cytometry, histological staining, and immunostaining, which collectively provided a detailed understanding of RBCs dynamics under HH conditions. In vitro analyses substantiated the decline in primary splenic macrophages and induction of ferroptosis under hypoxic treatment, effects that were relieved by pre-treatment with the ferroptosis inhibitor ferrostatin-1. Collectively, our data suggest that HH exposure initiates splenic ferroptosis, leading primarily to a decrease in RPMs. This decrease potentially impacts erythrophagocytosis, contributing to increased RBCs retention and subsequent splenomegaly. Such changes could potentially foster continuous RBCs production and accelerate the progress of HAPC. In conclusion, our study highlights the important role of the spleen and splenic macrophages in the pathogenesis of HAPC, providing crucial insights into the disease progression and potential therapeutic targets.

Removed in the revised version: This study investigates the role of the spleen and splenic macrophages in the development of high-altitude polycythemia (HAPC), a condition characterized by an imbalance between erythropoiesis and eryptosis imbalance, typically observed in high-altitude (HA) environments. We employed a mouse model subjected to hypobaric hypoxia (HH) to simulate the conditions of a 6000 m HA exposure. For in vitro examination, primary splenic macrophages were treated with 1% hypoxia. Our findings revealed that the HH treatment enhanced erythropoiesis, resulting in erythrocytosis, with marked splenic contraction evident, later progressing to splenomegaly over a 14-day period. HH exposure also impaired the ability of the spleen to process red blood cells (RBCs), predominantly due to a decrease in splenic macrophages within the red pulp. Furthermore, the application of HH treatment over 7 and 14-day intervals resulted in increased iron mobilization and onset of ferroptosis within the spleen, as corroborated by the expression of iron metabolism-related and ferroptosis-related proteins. The expression levels of these proteins mirrored gene expression levels in human peripheral blood mononuclear cells. Subsequent single-cell sequencing of the spleen demonstrated a substantial decrease in macrophages 7 days post-HH exposure. In vitro investigations confirmed the decline in primary splenic macrophages and induction of ferroptosis following hypoxic treatment, which were reversed by pre-treatment with the ferroptosis inhibitor ferrostatin-1. In summary, the data suggested that HH exposure instigates splenic ferroptosis, predominantly in the red pulp, thereby hampering the RBCs clearance in the spleen. This leads to increased 46 RBCs retention within the spleen, triggering splenomegaly, which may potentially foster continuous RBCs production and accelerate HAPC progression. The major conclusion from this study elucidates the critical role of spleen and splenic macrophages in the pathogenesis of HAPC.

Background

Earlier version: A plateau is a special environment characterized by low atmospheric pressure and low partial oxygen pressure. Long-term exposure to plateau environments may lead to chronic mountain disease [1]. High-altitude polycythemia (HAPC) is a common and widespread chronic mountain sickness characterized by excessive erythrocytosis [2]. Hypobaric hypoxia (HH) is the main cause of erythrocytosis, which in turn alleviates hypoxic conditions in tissues under high-altitude (HA) exposure [3]. In a healthy organism, a balance is maintained between erythropoiesis in the bone marrow and eryptosis in the spleen [4]. Even though HA/HH exposure can disrupt the equilibrium between RBC formation and clearance, healthy individuals can swiftly establish a new physiological RBC steady state [5]. However, in patients with HAPC, RBCs homeostasis is disrupted, failing to reach a state of equilibrium, which ultimately leads to a persistent increase in RBCs [4]. It is well-established that exposure to HA/HH can induce erythropoiesis, yet the pathogenesis of eryptosis under these conditions remains poorly understood.

Revised version: A plateau is a special environment characterized by low atmospheric pressure and low partial oxygen pressure. Long-term exposure to plateau environments may lead to chronic mountain disease [1]. High-altitude polycythemia (HAPC) is a common and widespread chronic mountain sickness characterized by excessive erythrocytosis [2]. Hypobaric hypoxia (HH) is the main cause of erythrocytosis, which in turn alleviates hypoxic conditions in tissues under high-altitude (HA) exposure [3]. In a healthy organism, a balance is maintained between erythropoiesis in the bone marrow (BM) and erythrophagocytosis in the spleen [4]. Even though HA/HH exposure can disrupt the equilibrium between RBCs formation and clearance, healthy individuals can swiftly establish a new steady state of RBCs homeostasis in chronic hypoxia [5]. This adaptation is a critical response to the altered oxygen availability characteristic of HA environments. However, in patients with HAPC, RBCs homeostasis is disrupted, failing to reach a state of equilibrium, which ultimately leads to a persistent increase in RBCs count [4]. This deviation from the normative adaptation process implies a pathological deviation from the usual compensatory mechanisms employed under chronic hypoxia conditions [6]. It is well-established that exposure to HA/HH can induce erythropoiesis, yet the pathogenesis of erythrophagocytosis under these conditions remains poorly understood [7].

Earlier version: The spleen plays a crucial role in maintaining erythropoietic homeostasis by effectively clearing impaired and senescent RBCs from circulation [6]. Particularly, red pulp macrophages (RPMs) within the spleen, serving as primary phagocytes, are responsible for clearing senescent, damaged, and abnormal erythrocytes from circulation to recycle iron [7]. RPMs initiate the process of RBC endocytosis and lysosomal digestion into heme, following the recognition of the RBCs earmarked for removal via their cell surface signals [8, 9]. Subsequently, heme oxygenase-1 (HO-1) decomposes the heme into biliverdin, carbon monoxide, and iron [10]. Most of the iron recycled from heme is transported out of the cell through a protein called ferroportin (Fpn). This iron then binds to transferrin (Tf), a plasma protein that transports iron in the blood. The iron-transferrin complex interacts with the transferrin receptor (TfR) on the cell membrane, contributing to the formation of RBCs in the erythroid compartment [11]. A portion of the released iron is loaded into cellular ferritin (Ft; including Ft-H and Ft-L). In the presence of oxygen, the Ft-H facilitates the oxidation of Fe2+ (ferrous ion) to Fe3+ (ferric ion). Subsequently, the Ft-L stores this Fe3+ [12]. The Ft-L features a nucleation site, consisting of a cluster of cavity-exposed carboxyl residues, which readily bind to Fe3+, thereby simplifying the storage process [13]. When the body’s iron metabolism is vigorous, nuclear receptor coactivator 4 (NCOA4) can recognize Ft-H, bring Ft into the lysosomal pathway for degradation, and release iron in Ft for iron utilization in the body [14]. However, some iron is free in cells in the form of ferrous iron (Fe2+), which is well known to be an important initiator of free radical oxidation. Moreover, the accumulation of large amounts of Fe2+ ions in cells may cause ferroptosis [15].

Revised version: The spleen plays a crucial role in maintaining erythropoietic homeostasis by effectively clearing impaired and senescent RBCs from circulation [8]. Particularly, red pulp macrophages (RPMs) within the spleen, serving as primary phagocytes, are responsible for clearing senescent, damaged, and abnormal erythrocytes from circulation to recycle iron [9]. RPMs initiate the process of RBCs endocytosis and lysosomal digestion into heme, following the recognition of the RBCs earmarked for removal via their cell surface signals [10, 11]. Subsequently, heme oxygenase-1 (HO-1) decomposes the heme into biliverdin, carbon monoxide, and iron [12]. Most of the iron recycled from heme is transported out of the cell through a protein called ferroportin (Fpn). This iron then binds to transferrin (Tf), a plasma protein that transports iron in the blood. The iron-transferrin complex interacts with the transferrin receptor (TfR) on the cell membrane, contributing to the formation of RBCs in the erythroid compartment [13]. A portion of the released iron is loaded into cellular ferritin (Ft; including Ft-H and Ft-L). In the presence of oxygen, the Ft-H facilitates the oxidation of Fe2+ (ferrous ion) to Fe3+ (ferric ion). Subsequently, the Ft-L stores this Fe3+ [14]. The Ft-L features a nucleation site, consisting of a cluster of cavity-exposed carboxyl residues, which readily bind to Fe3+, thereby simplifying the storage process [15]. When the iron metabolism is vigorous, nuclear receptor coactivator 4 (NCOA4) can recognize Ft-H, bring Ft into the lysosomal pathway for degradation, and release iron in Ft for iron utilization in the body [16]. However, some iron exists in cells in the form of ferrous iron (Fe2+), which is well known to be an important initiator of free radical oxidation [6]. Moreover, the accumulation of large amounts of Fe2+ in cells may cause ferroptosis [17].

Earlier version: Despite the extensive study of erythropoiesis under HA/HH conditions, the precise effects of HA/HH on erythrophagocytosis within the spleen remain largely unexplored. Considering the significant role of the spleen in RBC processing and the crucial function of RPMs in iron recycling from RBC clearance, we evaluated the impacts of HA exposure on the spleen/splenic macrophages using an HH-exposed mouse model (simulating 6000 m exposure conditions). In the current study, we sought to further investigate whether high-altitude/hypobaric hypoxia (HA/HH) can influence the erythrocyte disposal and iron recycling processes in macrophages. More specifically, we aimed to determine the potential impact of HA/HH exposure on the integrity of the erythrophagocytosis process. We discovered that exposure to HH triggered ferroptosis in the spleen, particularly in macrophages. This led to a reduction in macrophage numbers, which was subsequently followed by disruptions in eryptosis and iron recycling within the spleen. These findings may hold clinical significance, particularly in the context of continuous pathological erythrocytosis and the progression of HAPC under HA exposure.

Revised version: Despite the extensive study of erythropoiesis under HA/HH conditions, the precise effects of HA/HH on erythrophagocytosis within the spleen remain largely unexplored. Considering the significant role of the spleen in RBC processing and the crucial function of RPMs in iron recycling from RBC clearance, we evaluated the impacts of HA exposure on the spleen/splenic macrophages using an HH-exposed mouse model (simulating 6000 m exposure conditions). In the current study, we sought to further investigate whether HA/HH can influence the erythrocyte disposal and iron recycling processes in macrophages. More specifically, we aimed to determine the potential impact of HA/HH exposure on the integrity of the erythrophagocytosis process. We discovered that exposure to HH triggered ferroptosis in the spleen, particularly in macrophages. This led to a reduction in macrophage numbers, which was subsequently followed by disruptions in erythrophagocytosis and iron recycling within the spleen. These findings may hold clinical significance, particularly in the context of continuous pathological erythrocytosis and the progression of HAPC under HA exposure.

Methods

HH mouse exposure model

Earlier version: Male C57BL/6 male mice (at 8 weeks of age) were obtained from the animal experimental centre of Nantong University. Mice were adapted to the facilities for 3 days before experiments. Mouse models of HH exposure were established by placing the mice in the decompression chamber (TOW-INT TECH, ProOx-810, China) for 1 d, 2 d, 3 d, 7 d, and 14 d. The parameters were set as follows: 6000 m above sea level, oxygen partial pressure of 11 kPa, lifting speed of 20 m/s, chamber pressure of 54 kPa, temperature of 25 ℃, humidity of 40%, and 12-h light/dark cycles. Tuftsin (1.5 mg/kg, MCE, HY-P0240, USA) was used to stimulate the phagocytosis of macrophages [16], and was injected at 7 d and 11 d during the 14-d period of mice exposed to HH.

Revised version: Male C57BL/6 male mice (at 8 weeks of age) were obtained from the animal experimental centre of Nantong University. Mice were adapted to the facilities for 3 days before experiments. Mouse models of HH exposure were established by placing the mice in the decompression chamber (TOW-INT TECH, ProOx-810, China) for 1, 2, 3, 7, and 14 days. The parameters were set as follows: 6000 m above sea level, oxygen partial pressure of 11 kPa, lifting speed of 20 m/s, chamber pressure of 54 kPa, temperature of 25 ℃, humidity of 40%, and 12-h light/dark cycles. The hypobaric chamber was configured to ensure a ventilation rate of 25 air exchanges per minute. Tuftsin (0.15 mg/kg, MCE, HY-P0240, USA) was used to stimulate the phagocytosis of macrophages [18], and was injected at 7 and 11 days during the 14-days period of mice exposed to HH. The control group animals were maintained in a similar chamber for the same amount of time under normobaric normoxia (NN) conditions.

Added in the revised version: Splenectomy

Added in the revised version: After anesthesia, the mice were placed in a supine position, and their limbs were fixed. The abdominal operation area was skinned, disinfected, and covered with a sterile towel. A median incision was made in the upper abdomen, followed by laparotomy to locate the spleen. The spleen was then carefully pulled out through the incision. The arterial and venous directions in the splenic pedicle were examined, and two vascular forceps were used to clamp all the tissue in the main cadre of blood vessels below the splenic portal. The splenic pedicle was cut between the forceps to remove the spleen. The end of the proximal hepatic artery was clamped with a vascular clamp, and double or through ligation was performed to secure the site. The abdominal cavity was then cleaned to ensure there was no bleeding at the ligation site, and the incision was closed. Post-operatively, the animals were housed individually. Generally, they were able to feed themselves after recovering from anesthesia and did not require special care.

Splenic macrophage culture and treatments

After isoflurane anaesthesia, mice were transcardially perfused with precooled PBS, and then splenic tissue was collected. The spleen was then dissected into minuscule fragments, each approximately 1 mm3 in volume, and separated in cold, phenol red-free Hanks’ balanced salt solution (HBSS; Gibco, Grand Island, NY, USA). The cell extracts were centrifuged at 500 rpm for 5 min at 4 ℃. The pellets were then resuspended and incubated with ammonium-chloride-potassium (ACK) lysing buffer (Thermo Fisher Scientific, USA) for 3 min at room temperature. Next, cell extracts were purified in a four-phase Percoll (Sigma‒Aldrich, St. Louis, MO, USA) gradient: 0, 30, 40, and 50% in HBSS. After centrifugation at 500 rpm for 20 min at 4°C, the macrophage-enriched fraction was collected at the interface between the 30% and 40% phases. Cells were washed in two rounds with HBSS and then resuspended in AIM V medium (Thermo Fisher Scientific). The cells in this medium were cultured at 37 ℃ and 5% CO2 for 12 h. After the cells were treated with ferrostatin-1 (Fer-1, 2 μm; Selleck Chemicals, S7243, USA) 1 h in advance, the cells were placed in a hypoxia workstation (Invivo2, Ruskinn, UK) for 24 h.

Blood smear and hematological indices

Earlier version: The collected blood was gently mixed with EDTA (15 g/L) to obtain anticoagulant blood, which was detected by a hematology analyser (XE-5000, Sysmex Corporation, Japan). Blood smear preparation method: a blood sample is dropped onto a slide, and another slide is used to push the blood sample into a uniform thin layer at a constant speed. After the blood samples were dried, Wright’s stain solution was added and incubated for 1 min. Then, the slide was rinsed with water. Blood smears were photographed with a DM4000B microscope (Leica, Germany).

Revised version: The collected blood was gently mixed with EDTA (15 g/L) to obtain anticoagulant blood, which was detected by a hematology analyser (XE-5000, Sysmex Corporation, Japan). Blood smear preparation method: a blood sample is dropped onto a slide, and another slide is used to push the blood sample into a uniform thin layer at a constant speed. After the blood samples were dried, Wright stain solution (G1040, Solarbio, China) was added and incubated for 1 min. Then, the slide was rinsed with water. Blood smears were photographed with a DM4000B microscope (Leica, Germany).

Added in the revised version: Lillie staining and Hematoxylin and Eosin (HE) staining

Removed in the revised version: Lillie staining

Earlier version: The Lillie staining method was employed to ascertain Fe2+ deposition within the spleen, according to the protocol delineated by Liu et al [17]. Briefly, paraffin-embedded spleen sections underwent a process of dewaxing and rehydration via xylene and gradient alcohol. Then, the Lillie staining solution (G3320, Solarbio, China) was applied at 37°C for 50 min in a light-restricted environment, improved by a 5-min PBS wash. Subsequently, a mixture comprising 30 % hydrogen peroxide and methanol was incubated for 20 min, and then the sections were subjected to two 5-min washes with PBS. The DAB developing solution, which is not included in the Lillie bivalent iron staining kit, was applied for light-sensitive staining. The sections were stained with a nuclear fast red solution in a light-restricted setting for 10 minutes, followed by a 5-second wash with ddH2O. Lastly, the sections were dehydrated with gradient alcohol, cleared with xylene, and mounted with neutral resin.

Revised version: The Lillie staining method was employed to ascertain Fe2+ deposition within the spleen, according to the protocol delineated by Liu et al [19]. Briefly, paraffin-embedded spleen sections underwent a process of dewaxing and rehydration via xylene and gradient alcohol. Then, the Lillie staining solution (G3320, Solarbio, China) was applied at 37°C for 50 min in a light-restricted environment, improved by a 5-min PBS wash. Subsequently, a mixture comprising 30% hydrogen peroxide and methanol was incubated for 20 min, and then the sections were subjected to two 5-min washes with PBS. The DAB developing solution, which is not included in the Lillie bivalent iron staining kit, was applied for light-sensitive staining. The sections were stained with a nuclear fast red solution in a light-restricted setting for 10 min, followed by a 5-second wash with ddH2O. Lastly, the sections were dehydrated with gradient alcohol, cleared with xylene, and mounted with neutral resin. HE staining was employed in the spleen according to our previous study [20].

Added in the revised version: Wright-Giemsa composite stain

Removed in the revised version: Wright staining

Earlier version: Mice were anaesthetized and perfused with saline. The spleen was subsequently excised to prepare a single cell suspension, achieved by filtering through a 70 μm filter. Samples were centrifuged at 500 g for 5 min for cell separation. The cell pellet was subjected to three PBS washes. The suspension was then smeared onto a slide and air-dried. Following this, the slide was stained with Wright solution (G1040, Solarbio, China) for 3 min. An equivalent volume of pH 6.4 phosphate buffer was introduced, and the slide was gently agitated to allow mixing with the Wright’s stain solution for 5 min. After washing and drying, a microscope (DM4000B, Leica, Germany) was utilized for visualization and image acquisition.

Revised version: Mice were anaesthetized and perfused with saline. The spleen was subsequently excised to prepare a single cell suspension, achieved by filtering through a 70 µm filter. Samples were centrifuged at 500 g for 5 min for cell separation. The cell pellet was subjected to three PBS washes. The suspension was then smeared onto a slide and air-dried. Following this, the slide was stained with Wright-Giemsa composite stain (G1020, Solarbio, China) for 3 min. An equivalent volume of pH 6.4 phosphate buffer was introduced, and the slide was gently agitated to allow mixing with the Wright’s stain solution for 5 min. After washing and drying, a microscope (DM4000B, Leica, Germany) was utilized for visualization and image acquisition.

Western blot

Earlier version: Tissue and cell samples were homogenized on ice after cell lysates and protease inhibitors were added. Centrifugation was performed at 12,000 rpm for 15 min at 4°C to collect the supernatant. Protein concentration was determined by the BCA detection method. A sample containing 30 μg protein was loaded and run in each well of SDS‒ PAGE gels. The membranes were incubated with the following antibodies: HO-1 (1:2000, Abcam, ab13243, USA); Ft-L (1:1000, Abcam, ab69090, USA); Ft-H (1:1000, Novus, NBP1-31944, USA); NCOA4 (1:1000, Santa Cruz, sc-373739, USA); TfR (1:1000, Thermo Fisher, 13-6800, USA); Fpn (1:1000, Novus, NBP1-21502, USA); ACSL4 (1:1000, Santa Cruz, sc-271800, USA); xCT (1:1000, Proteintech, 26864-1-AP, USA); Gpx4 (1:1000, Abcam, ab125066, USA); CD206 (1:1000, RD, AF2535, USA); CD16 (1:2000, RD, AF1960, USA); and β-actin (1:1000, Sigma‒Aldrich, A5316, USA). The secondary antibodies were as follows: goat anti-mouse (1:10000, Jackson, USA) and goat anti-rabbit (1:10000, Jackson, USA). The gray value of specific blots was scanned and analysed using ImageJ software (National Institute of Health, USA).

Revised version: Tissue and cell samples were homogenized on ice after cell lysates and protease inhibitors were added. Centrifugation was performed at 12,000 rpm for 15 min at 4°C to collect the supernatant. Protein concentration was determined by the BCA detection method. A sample containing 30 μg protein was loaded and run in each well of SDS‒ PAGE gels. The membranes were incubated with the following antibodies: HO-1 (1:2000, Abcam, ab13243, USA); Ft-L (1:1000, Abcam, ab69090, USA); Ft-H (1:1000, Novus, NBP1-31944, USA); NCOA4 (1:1000, Santa Cruz, sc-373739, USA); TfR (1:1000, Thermo Fisher, 13-6800, USA); Fpn (1:1000, Novus, NBP1-21502, USA); ACSL4 (1:1000, Santa Cruz, sc-271800, USA); xCT (1:1000, Proteintech, 26864-1-AP, USA); Gpx4 (1:1000, Abcam, ab125066, USA); CD206 (1:1000, RD, AF2535, USA); and β-actin (1:1000, Sigma‒Aldrich, A5316, USA). The secondary antibodies were as follows: goat anti-mouse (1:10000, Jackson, USA) and goat anti-rabbit (1:10000, Jackson, USA). The gray value of specific blots was scanned and analysed using ImageJ software (National Institute of Health, USA).

RT‒PCR analysis

Earlier version: Total RNA extraction and RT‒PCR analysis were performed essentially as described previously [18]. Primer sequences were as follows:

Revised version: Total RNA extraction and RT‒PCR analysis were performed essentially as described previously [21]. Primer sequences were as follows:

GEO analysis

GSE46480 samples were found in the GEO database of NCBI with the keyword “High Altitude” search. The data set was divided into two groups, one for the plain (Base Camp) and the other for the plateau (Altitude), with a sample size of 98 for each group. Blood samples were collected from the same person at McMurdo Station (48 m) and immediately transferred to Amundsen-Scott South Pole Station (2835 m) on the third day. R language was used for data analysis, and GraphPad Prism was used for statistics.

Malondialdehyde (MDA), Cysteine (Cys) and Glutathione (GSH) content detection

MDA was detected by a Micro Malondialdehyde Assay Kit (Solarbio, BC0025, China). Cys was detected by a Micro Cysteine Assay Kit (Solarbio, BC0185, China). GSH was detected by a Micro Reduced Glutathione Assay Kit (Solarbio, BC1175, China).

Immunofluorescence

Earlier version: Spleen sections of 9 μm thickness were flash-frozen and stored at –80°C. Before immunofluorescence staining, these sections were allowed to acclimate to room temperature for 30 minutes, then incubated in 10% BSA-PBS for 10 minutes. Next, the sections were incubated overnight with F4/80-PE (1:200, Biolegend, 123110, USA) and then washed thrice with PBS. Due to the limitations of F4/80 as a definitive macrophage marker, we concurrently conducted immunohistochemical analysis for heme oxygenase-1 (HO-1) (1:200, Abcam, ab13243, USA). Following this, sections were washed thrice with 0.05% PBST, incubated with secondary antibodies DAR-Alexa555 at room temperature for 2 hours, washed thrice with PBS again, and finally sealed with 50% glycerine-PBS. Fluorescence microscopy images were captured using a confocal laser scanning microscope (SP8, Leica Microsystems, Wetzlar, Germany). To quantify the immunostaining intensity of F4/80 and HO-1, we utilized ImageJ software. Specifically, we captured multiple fields of view for each slide, and the software was used to measure the mean intensity of the fluorescent signal in these images. These measurements were then normalized to the NN group to account for any experimental variations.

Revised version: Spleen sections of 9 μm thickness were flash-frozen and stored at -80°C. Before immunofluorescence staining, these sections were allowed to acclimate to room temperature for 30 min, then incubated in 10% BSA-PBS for 10 min. Next, the sections were incubated overnight with F4/80-PE (1:200, Biolegend, 123110, USA) or F4/80 (1:100, ab16911, Abcam, USA) and then washed thrice with PBS. Due to the limitations of F4/80 as a definitive macrophage marker, we concurrently conducted immunohistochemical analysis for heme oxygenase-1 (HO-1) (1:200, Abcam, ab13243, USA), CD11b (1:100, 14-0112-82, Thermo, USA) and CD68 (1:100, ab125212, Abcam, USA), respectively. In the RBCs detection experiment, Band3 primary antibody (1:100, 2813101-AP, Proteintech, China) and PKH67 green fluorescent cell linker (1:250, MIDI67, Sigma, USA) were used. Following this, sections were washed thrice with 0.05% PBST, incubated with secondary antibodies DAR-Alexa555 at room temperature for 2 h, washed thrice with PBS again, and finally sealed with 50% glycerine-PBS. Fluorescence microscopy images were captured using a confocal laser scanning microscope (SP8, Leica Microsystems, Wetzlar, Germany). To quantify the immunostaining intensity of F4/80, HO-1, CD11b, CD68, we utilized ImageJ software. Specifically, we captured multiple fields of view for each slide, and the software was used to measure the mean intensity of the fluorescent signal in these images. These measurements were then normalized to the NN group to account for any experimental variations.

Tissue iron staining (DAB-enhanced Perls’ staining)

Earlier version: DAB-enhanced Perls’ staining for the spleen paraffin sections was performed as described previously [19]. Briefly, the sections were washed with PBS and cultured in freshly prepared Perls’ solution (1% potassium ferricyanide in 0.1 M hydrochloric acid buffer). The slides were then immersed and stained with DAB. All slides were counterstained with hematoxylin and visualized under a DM4000B microscope (Leica, Germany). Data were collected from three fields of view per mouse and semiquantitatively analysed with ImageJ software. Quantitative results of iron staining were finally normalized to the NN control group.

Revised version: DAB-enhanced Perls’ staining for the spleen paraffin sections was performed as described previously [20]. Briefly, the sections were washed with PBS and cultured in freshly prepared Perls’ solution (1% potassium ferricyanide in 0.1 M hydrochloric acid buffer). The slides were then immersed and stained with DAB. All slides were counterstained with hematoxylin and visualized under a DM4000B microscope (Leica, Germany). Data were collected from three fields of view per mouse and semiquantitatively analysed with ImageJ software. Quantitative results of iron staining were finally normalized to the NN control group.

Flow cytometry

Earlier version: The steps of reticulocyte ratio detection were as follows: Whole blood was extracted from the mice and collected into an anticoagulant tube, which was then set aside for subsequent thiazole orange (TO) staining [20]. The experimental tube and negative control tube were prepared, 125 μL normal saline, 4 μL anticoagulant and 125 μL TO working solution (1 μg/mL, Sigma, 390062, USA) were added to each tube, and 250 μL normal saline and 4 μL anticoagulant were added to each tube of the negative control tube. After incubation at room temperature for 1 h, flow cytometry analysis was carried out by using the FL1 (488 nm/525 nm) channel.

Revised version: The steps of reticulocyte ratio detection were as follows: Whole blood was extracted from the mice and collected into an anticoagulant tube, which was then set aside for subsequent thiazole orange (TO) staining [22]. The experimental tube and negative control tube were prepared, 125 μL normal saline, 4 μL anticoagulant and 125 μL TO working solution (1 μg/mL, Sigma, 390062, USA) were added to each tube, and 250 μL normal saline and 4 μL anticoagulant were added to each tube of the negative control tube. After incubation at room temperature for 1 h, flow cytometry analysis was carried out by using the FL1 (488 nm/525 nm) channel.

The steps for the determination of intracellular divalent iron content and lipid peroxidation level were as follows: Splenic tissue was procured from the mice and subsequently processed into a single-cell suspension using a 40 μm filter. The RBCs within the entire sample were subsequently lysed and eliminated, and the remaining cell suspension was resuspended in PBS in preparation for ensuing analyses. A total of 1 × 106 cells were incubated with 100 μL of BioTracker Far-red Labile Fe2+ Dye (1 mM, Sigma, SCT037, USA) for 1 h or C11-BODIPY 581/591 (10 μM, Thermo Fisher, D3861, USA) for 30 min. After the cells were washed with PBS twice, flow cytometry analysis was carried out by using the FL6 (638 nm/660 nm) channel for determination of intracellular divalent iron content or the FL1 (488 nm/525 nm) channel for determination of lipid peroxidation level.

Earlier version: The steps for detecting the mortality of spleen cells and peritoneal macrophages were as follows: 1 × 106 cells were incubated at room temperature for 40 min with 2 μM Calcein AM and 8 μM PI. Flow cytometry analysis was carried out by using FL1 (488 nm/525 nm, Calcein AM) and FL3 (488 nm/620 nm, PI) channels.

Revised version: The steps for detecting the mortality of spleen cells and macrophages were as follows: 1 × 106 cells were incubated at room temperature for 40 min with 2 μM Calcein AM and 8 μM Propidium Iodide (PI). Flow cytometry analysis was carried out by using FL1 (488 nm/525 nm, Calcein AM) and FL3 (488 nm/620 nm, PI) channels.

Earlier version: The steps for detecting the number of M1/M2 macrophages in the spleen were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE (1:200), CD86-PE/Cyanine7 (1:20, Biolegend, 105014, USA), and CD206-Alexa 647 (1:80, BD, 565250) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80-PE), FL5 (488 nm/755 nm, CD86-PE/Cyanine7) and FL6 (638 nm/660 nm, CD206-Alexa 647) channels.

Revised version: The steps for detecting the number of RPMs in the spleen were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE (1:1000, 565410, BD, USA) and CD11b PE-CY7 M1/70 (1:1000, 552850, BD, USA) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80- PE) and FL5 (488 nm/755 nm, CD11b PE-CY7 M1/70) channels.

The steps for detecting the number of monocytes in blood, spleen and bone marrow were as follows: 1 × 106 cells were incubated with 2% mouse serum-PBS for 10 min, incubated with F4/80-PE, CD11b-PE/CY7 (1:2000, BD, 552850, USA), and Ly6C-APC (1:2000, Thermo Fisher, 17-5932-82) for 30 min, and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL2 (488 nm/575 nm, F4/80- PE), FL5 (488 nm/755 nm, CD11b-PE/CY7) and FL6 (638 nm/660 nm, Ly6C-APC) channels.

Added in the revised version: The steps employed for assessing erythrophagocytosis in the spleen was scrupulously performed as follows: Initially, RBCs were harvested from mice and subjected to in vitro labelling with PKH67 cell linker. Subsequently, these PKH67-labeled RBCs were administered into mice, followed by exposure to either NN or HH. After a duration of 7 or 14 days of exposure, splenic single-cell suspensions were carefully prepared. From these suspensions, a quantity of 1 × 106 cells was isolated and washed in PBS for a period of 5 min. The analytical phase involved flow cytometry, utilizing the FL1 channel (490 nm excitation/502 nm emission) specifically used for detecting PKH67. This approach facilitated the precise quantification and analysis of erythrophagocytosis within the spleen under the specified experimental conditions.

Phagocytosis of E. coli and RBCs

E. coli was labelled with Cy5.5 (5 mg/mL), and RBCs were labelled with NHS-biotin (20 mg/mL). Macrophages (1 × 106) were coincubated with E. coli-Cy5.5 or RBC-Biotin for 30 min and washed with 2% mouse serum-PBS twice. Flow cytometry analysis was carried out by using FL6 (638 nm/660 nm) to determine the phagocytosis of spleen by E. coli. Macrophages (coincubated with Biotin-RBCs) were incubated with streptavidin-FITC for 3 h and washed twice with 2% mouse serum PBS. FL1 (488 nm/525 nm) was used for flow cytometry analysis to determine the phagocytosis of RBCs in the spleen.

Single-cell RNA Sequencing

The spleen tissues were surgically removed and stored in MACS Tissue Storage Solution (Miltenyi Biotec, Bergisch Gladbach, Germany) until processing. Single-cell dissociation was performed by the experimentalists at the GENECHEM laboratory (Shanghai, China). Dissociated single cells were then stained for viability assessment using Calcein-AM (BD Biosciences, USA) and Draq7 (BD Biosciences). The BD Rhapsody system was used to capture transcriptomic information from single cells. Single-cell capture was achieved by random distribution of a single-cell suspension across >200,000 microwells through a limited dilution approach. Beads with oligonucleotide barcodes were added to saturation so that a bead was paired with a cell in a microwell. The cells were lysed in the microwell to hybridize mRNA molecules to barcoded capture oligos on the beads. Beads were collected into a single tube for reverse transcription and ExoI digestion. Upon cDNA synthesis, each cDNA molecule was tagged on the 5′ end (that is, the 3′ end of an mRNA transcript) with a unique molecular identifier (UMI) and cell barcode indicating its cell of origin. Whole transcriptome libraries were prepared using the BD Rhapsody single-cell whole-transcriptome amplification (WTA) workflow, including random priming and extension (RPE), RPE amplification PCR and WTA index PCR. The libraries were quantified using a High Sensitivity D1000 ScreenTape (Agilent) and High Sensitivity D1000 Reagents (Agilent) on a 4150 TapeStation System (Agilent, Palo Alto, CA, USA) and the Qubit High Sensitivity DNA assay (Thermo Fisher Scientific). Sequencing was performed on an Illumina sequencer (Illumina Nova Seq 6000, San Diego, CA) in a 150 bp paired-end run.

Statistical analysis

Statistical analysis was conducted using GraphPad Prism 8.0 software. All values were represented as the mean ± standard error (SEM). The homogeneity of variance was verified before the application of a parametric test. A two-tailed Student’s t-test was employed for data from two groups, while a one-way analysis of variance (ANOVA) with multiple comparisons using Tukey’s post hoc test was utilized for data from multiple groups. A p-value of less than 0.05 was considered statistically significant.

Results

HH exposure promotes erythrocytosis in mice

Earlier version: To mimic 6000 m HA exposure, we placed C58BL/6 mice in an animal hypobaric oxygen chamber and detected the blood indices in the blood of the mice after HH exposure for different times. The blood smear showed that the number of RBCs was increased from 3 to 14 days after HH exposure compared with the normobaric normoxia (NN) group (Figure 1A). Routine blood tests further confirmed the results in blood smears, which showed that the RBC number (Figure 1B), HGB content (Figure 1C) and HCT value (Figure 1D) were all increased significantly to varying degrees, while MCH was not changed after HH exposure (Figure 1E). We further performed flow cytometry by TO staining to detect reticulocytes after 7 and 14 days of HH treatment. The results showed that at both HH 7 and 14 d of exposure, the reticulocyte proportion was increased (Figure 1F). These results suggested that the HH-treated mouse model mimics HA exposure well, which promotes erythropoiesis and results in erythrocytosis in mice following HH exposure.

Revised version: To mimic 6000 m HA exposure, we placed C58BL/6 mice in an animal hypobaric oxygen chamber and detected the blood indices in the blood of the mice after HH exposure for different times. The blood smear showed that the number of RBCs was increased from 3 to 14 days after HH exposure compared with the NN group (Figure 1A). Routine blood tests further confirmed the results in blood smears, which showed that the RBC number (Figure 1B), HGB content (Figure 1C) and HCT value (Figure 1D) were all increased significantly to varying degrees, while MCH was not changed after HH exposure (Figure 1E). We further performed flow cytometry using TO staining to detect reticulocytes after 3, 7, and 14 days of HH treatment. The results showed that the proportion of reticulocytes increased after 7 and 14 days of HH exposure (Figure 1, F-G), and the number of RBCs reached the peak after 7 days of HH exposure. These results suggested that the HH-treated mouse model effectively mimics HA exposure, which promotes erythropoiesis and results in erythrocytosis in mice following HH exposure.

Earlier version: HH exposure promotes erythrocytosis in mice.

Revised version: HH exposure promotes the induction of erythrocytosis in mice.

Earlier version: The C57BL/6 mice were treated with normobaric normoxia (NN) and hypobaric hypoxia (HH) for 3, 7 and 14 days, and blood was collected for the following tests. (A) Analysis of the number and morphology of RBCs via a blood cell smear. Routine blood analysis was performed for RBC counts (B), as well as HGB (C), HCT (D), and MCH content (E). (F) Representative dot plots (from three experiments) indicate TO-positive cells (reticulocytes) in whole blood. Data are expressed as means ± SEM (n = 5 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: C57BL/6 mice were subjected to either normobaric normoxia (NN) or hypobaric hypoxia (HH) conditions for durations of 3, 7, and 14 days. After these treatments, blood samples were collected for a comprehensive analysis. (A) Morphological evaluation of RBCs was conducted via blood smear examination (Wright staining). Routine hematological assessments were performed, encompassing RBCs counts (B), hemoglobin (HGB) levels (C), hematocrit (HCT) percentages (D), and mean corpuscular hemoglobin (MCH) content (E). (F) Flow cytometric analysis was employed to identify TO-positive cells, indicative of reticulocytes, in the whole blood samples. (G) The proportions of TO-positive RBCs in the blood were depicted in bar graphs. The data are presented as means ± SEM for each group (n = 5 per group). Statistical significance is denoted by * P < 0.05, ** P < 0.01, *** P < 0.001, relative to the NN group or as specified.

Spleen inhibits the immoderate increase in RBCs under HH conditions

Earlier version: To determine the roles of the spleen in RBC homeostasis under HA/HH, we investigated the effects of HH on the morphology, volume and weight of the spleen as well as erythrocyte indices. As shown in Figure 2, the spleen volume and weight were decreased significantly after HH exposure for 1 day compared to NN treatment (Figure 2, A-C). However, the spleen was obviously enlarged from 2 to 14 days after HH exposure (Figure 2, A-C). The results indicated that the spleen contracted, the stored RBCs in the spleen were released into the blood at 1 day, and the RBCs were produced and/or retained in the spleen from 2 to 14 days after HH exposure. In our research, we also examined the influence of the spleen on RBCs homeostasis under HH conditions. We investigated whether the role of spleen in RBCs clearance under HH conditions could be compensated by the liver or other components of the mononuclear macrophage system. To conduct this, we performed splenectomies on mice and subsequently exposed them to HH conditions for 14 days. This allowed us to monitor RBCs counts and blood deposition. Our findings indicated that, in comparison to both the splenectomized mice under NN conditions and the sham-operated mice exposed to HH, erythrocyte deposition (Figure 2D) and counts (Figure 2E), as well as HGB (Figure 2F) and HCT (Figure 2G) levels, significantly increased 14 days post-splenectomy under HH conditions. Meanwhile, MCH levels remained stable (Figure 2H). These indices did not vary in mice, regardless of whether they had undergone a splenectomy, under NN conditions (Figure 2D and E). These results suggest that the liver and other mononuclear macrophages do not effectively compensate under HH conditions. Furthermore, our research highlights the crucial role of the spleen in maintaining RBCs homeostasis under conditions of HH. The study suggests that the spleen functions to regulate the proliferation of RBCs, preventing their overproduction in response to HH exposure.

Revised version: To determine the roles of the spleen in RBC homeostasis under HA/HH, we investigated the effects of HH on the morphology, volume, and weight of the spleen as well as erythrocyte indices. As shown in Figure 2, the spleen volume and weight were decreased significantly after HH exposure for 1 day compared to NN treatment (Figure 2, A-C). However, the spleen was obviously enlarged from 2 to 14 days after HH exposure (Figure 2, A-C). The results indicated that the spleen contracted, the stored RBCs in the spleen were released into the blood at 1 day, and the RBCs were produced and/or retained in the spleen from 2 to 14 days after HH exposure. In our research, we also examined the influence of the spleen on RBCs homeostasis under HH conditions. We investigated whether the role of spleen in RBCs clearance under HH conditions could be compensated by the liver or other components of the mononuclear macrophage system. To conduct this, we performed splenectomies on mice and subsequently exposed them to HH conditions for 14 days. This allowed us to monitor RBCs counts and blood deposition. Our findings indicated that, in comparison to both the splenectomized mice under NN conditions and the sham-operated mice exposed to HH, erythrocyte deposition (Figure 2D) and counts (Figure 2E), as well as HGB (Figure 2F) and HCT (Figure 2G) levels, significantly increased 14 days post-splenectomy under HH conditions. Meanwhile, MCH levels remained stable (Figure 2H). These indices did not vary in mice, regardless of whether they had undergone a splenectomy, under NN conditions (Figure 2D and E). These results indicate that in the splenectomized group under NN conditions, erythrophagocytosis is substantially compensated for by functional macrophages in other tissues. However, under HH conditions, our data also suggest that the spleen plays an important role in managing erythrocyte turnover, as indicated by the significant impact of splenectomy on erythrophagocytosis and subsequent RBCs dynamics.

The spleen plays an important role in suppressing the immoderate increase in RBCs under HH conditions.

Earlier version: C57BL/6 mice with or without splenectomy were treated with NN and HH for varying durations, and the spleen and blood were collected for subsequent analyses. (A) Morphological observation, (B) spleen volume, and (C) spleen weight were determined post-HH exposure. (D-H) Blood observation and hematological index detection followed. Data are expressed as the means ± SEM (n = 9 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: C57BL/6 mice with or without splenectomy were treated with NN and HH for varying durations, and the spleen and blood were collected for subsequent analyses. (A) Morphological observation, (B) Spleen volume, and (C) spleen weight was determined. (D-H) Blood observation and hematological index detection followed. Data are expressed as the means ± SEM (n = 9 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Earlier version: Reduction of splenic macrophages due to HH exposure

Revised version: HH exposure leads to a decrease in splenic macrophages

Added in the revised version: Considering macrophages as the primary cell type responsible for processing RBCs within the spleen under physiological conditions, we subsequently investigated the population and activity of these macrophages after exposure to HH for 7 or 14 days, employing flow cytometry and single-cell sequencing techniques. Calcein/PI double staining, examined via flow cytometry, revealed a significant decrease in viable splenic cells and a concomitant increase in dead cells following 7 and 14 days of HH exposure (Figure 3, A-C). This observation was further substantiated by single-cell sequencing, which elucidated a pronounced reduction in the population of splenic macrophages after 7 days of HH exposure (Figure 3, D-F). Additionally, flow cytometry results indicated a marked decrease in the population of RPMs, identified as F4/80hiCD11blo, in the spleen post 7 days of HH exposure (Figure 3G). Complementary to these findings, immunofluorescence detection demonstrated a significant diminution in RPMs, characterized by F4/80hiCD11blo and F4/80hiCD68hi cell populations, after both 7 and 14 days of HH exposure (Figure S1). Furthermore, our single-cell sequencing analysis of the spleen under NN conditions revealed a predominant association of RPMs with Cluster 0. Intriguingly, HH exposure led to a notable reduction in the abundance of RPMs within this cluster. Pseudo-time series analysis provided insights into the transitional dynamics of spleen RPMs, indicating a shift from Cluster 2 and Cluster 1 towards Cluster 0 under NN conditions. However, this pattern was altered under HH exposure, where a shift from Cluster 0 and Cluster 1 towards Cluster 2 was observed (Figure S2).

Removed in the revised version: Considering macrophages as the primary cell type responsible for processing RBCs within the spleen under physiological conditions, we subsequently investigated the population and activity of these macrophages after exposure to HH for 7 or 14 days, employing flow cytometry and single-cell sequencing techniques. Calcein/PI double staining, examined via flow cytometry, revealed a significant decrease in viable spleen cells and a concurrent increase in dead cells following 7 and 14 days of HH exposure (Figure 3, A-C). Single-cell sequencing further disclosed a significant reduction in splenic macrophages after 7 days of HH exposure (Figure 3, D-F). Additionally, the expressions of CD16 (M1 macrophage marker) and CD206 (M2 macrophage marker) were observed to decrease after both 7 and 14 days of HH exposure (Figure 3, G-L). Western blot analysis corroborated the flow cytometry findings, showing a decrease in mature macrophages (F4/80+/CD11b+), including both M1-type (F4/80+/CD86+) and M2-type (F4/80+/CD206+) macrophages, in the spleens of mice after 7 days of HH exposure (Figure 3M). To further elucidate the migration and differentiation of monocytes from the bone marrow to the spleen, we analyzed the expression of chemokines CCL2, CCL7, Csf1, and Csf2 in the spleen using qPCR and determined the number of monocytes in the bone marrow (BM) and spleen via flow cytometry. Our results indicated a significant reduction in the expression of CCL2, CCL7 (Figure 3N), Csf1, and Csf2 (Figure 3O) in the spleen after 7 and 14 days of HH exposure. Furthermore, the number of monocytes (Ly6C+/CD11b+) in the bone marrow (Figure 3, P-R) and spleen (Figure 3, P and S-T) also exhibited a decline after 7 and 14 days of HH exposure. 373 We evaluated the depletion of macrophages in the spleen under HH conditions by examining the expression and distribution of HO-1 and F4/80. Figures 3U-W depict a decrease in both HO-1 and F4/80, predominantly within the spleen’s red pulp following HH exposure. Together, these findings indicate a reduction in the number of splenic macrophages after HH exposure, which could impair the spleen’s capacity to process erythrocytes.

HH exposure results in a decrease in the number of splenic macrophages.

Earlier version: C57BL/6 mice were treated with NN and HH for 7 or 14 days, and the spleen and blood were collected for subsequent detection. (A) Calcein/PI double staining was analysed by flow cytometry. (B and C) The bar graphs represent cell death proportions in total splenic cells. (D) Uniform manifold approximation and projection (UMAP) visualized spleen cell clusters. Each color represents a cluster of specific spleen cells, identified by a distinct gene expression profile. (E) The proportions of different macrophage clusters derived from either NN-or HH-treated mouse spleens were analysed. (F) Individual macrophages derived from either NN-or HH-treated mouse spleens are denoted. (G and J) Western blot detecting CD16 and CD206 protein expression in the spleen. (H-I and K-L) Statistical analysis of CD16 and CD206 expression in H and K. (M) F4/80 staining with CD11b, CD86 or CD206 in splenic cells was analysed by flow cytometry. (N) qPCR analysis of CCL2 and CCL7 expression in the spleen. (O) qPCR analysis of Csf1 and Csf2 expression in the spleen. (P) Flow cytometry detection of CD11b and Ly6C double-stained cells in the BM and spleen. CD11bhiLy6Chi cell proportions in the BM (Q-R) and spleen (S-T) are given as bar graphs. (U) HO-1 and F4/80 expression in the spleen after 0 (NN), 7, and 14 914 days of HH exposure. (V) The relative fluorescence of HO-1 and (W) F4/80 915 expression in the spleen as described in (U) was quantified. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: C57BL/6 mice were treated with NN and HH for 7 or 14 days, and the spleen and blood were collected for subsequent analysis. (A) Calcein/PI double staining was analysed by flow cytometry to assess cell viability. (B and C) The bar graphs represent the proportions of cell death in total splenic cell population. (D) Uniform Manifold Approximation and Projection (UMAP) provided a visualization of spleen cell clusters, with each color representing a unique cluster characterized by specific gene expression profiles. (E) Comparative analysis of the proportions of distinct macrophage clusters in spleens from NN- and HH-treated mice. (F) Presentation of individual macrophages from spleens of mice treated with either NN or HH. (G) Analysis of F4/80 and CD11b expression in splenic cells, conducted via flow cytometry. (H) qPCR analysis evaluated CCL2 and CCL7 gene expression in the spleen. (I) qPCR analysis of Csf1 and Csf2 expression levels in the spleen. (J) Flow cytometry facilitated the detection of CD11b and Ly6C double-stained cells in BM and spleen. (K-L) Bar graphs represent the proportions of CD11bhiLy6Chi cells in the BM, while (M-N) delineate those in the spleen. (O) HO-1 and F4/80 expression levels in the spleen were monitored after 0 (NN), 7, and 14 days of HH exposure. (P) The relative fluorescence intensities of HO-1 and (Q) F4/80, as outlined in (O), were quantitatively assessed. Data are expressed as means ±SEM for each group (n = 3 per group). Statistical significance is indicated by * P < 0.05, ** P < 0.01, *** P < 0.001 when compared to the NN group or the indicated group.

Added in the revised version: To further elucidate the migration and differentiation of monocytes from the bone marrow to the spleen, we analysed the expression of chemokines CCL2, CCL7, Csf1, and Csf2 in the spleen using qPCR and determined the number of monocytes in the bone marrow (BM) and spleen via flow cytometry. Our results indicated a significant reduction in the expression of CCL2, CCL7 (Figure 3H), Csf1, and Csf2 (Figure 3I) in the spleen after 7 and 14 days of HH exposure. Furthermore, the number of monocytes (Ly6C+/CD11b+) in the bone marrow (Figure 3, J-L) and spleen (Figure 3, J and M-N) also exhibited a decline after 7 and 14 days of HH exposure. We evaluated the depletion of macrophages in the spleen under HH conditions by examining the expression and distribution of HO-1 and F4/80. There was a decrease in both HO-1 and F4/80, predominantly within the splenic red pulp following HH exposure (Figures 3, O-Q). Together, these findings indicate a reduction in the number of splenic macrophages after HH exposure, which could impair the spleen’s capacity to process erythrocytes.

Added in the revised version: HH exposure suppresses erythrophagocytosis of RPMs in spleen

Removed in the revised version: The impact of HH exposure on erythrophagocytosis and iron processing capacity of the spleen

Earlier version: We investigated the influence of HH exposure on erythrocyte phagocytosis and heme iron recycling within splenic macrophages. We administered NHS-biotin intravenously to the mice followed by HH treatment in vivo. The clearance of RBCs was monitored by detecting the retention of biotin-labelled RBCs using flow cytometry in both the blood and spleen of the HH-treated mice. We observed a substantial decrease in biotin-labelled RBCs in the blood and spleen cells at 7-and 14-days post HH exposure (Figure 4, A-D). Despite the natural decrease of labelled RBC over time, the decay was more pronounced in the HH exposure group compared to the normobaric normoxia (NN) group in both the blood (Figure 4, A and C) and spleen (Figure 4, B and D). The phagocytic ability of mouse spleen macrophages towards RBCs was significantly more impaired after HH exposure, particularly on the 14th day (Figure 4, A-D). Additionally, to substantiate our findings, we quantified the retention of RBC counts by Wright staining in splenic cells after 7 and 14 days of HH exposure (Figure 4, E-F). This confirmed impaired erythrophagocytosis in the spleen under HH exposure, as evidenced by an increase in deformation and retention of RBCs following 7 and 14 days of HH exposure. Furthermore, we employed Tuftsin to stimulate phagocytosis of macrophages and carried out immunofluorescence and iron staining to examine the heme iron recycling capability of the spleen under HH conditions (Figure 4G). As depicted in Figure 4H, compared with the NN group, F4/80 expression and iron deposition in the red pulp were significantly reduced after HH exposure. However, Tuftsin led to an induction of F4/80 expression and iron deposition in the red pulp of the spleen under HH (Figure 4, I-L). These results imply that the capacity for erythrocyte phagocytosis and heme iron recycling by the spleen was compromised, principally due to the reduction in macrophages in the red pulp under HH conditions.

Revised version: We investigated the influence of HH exposure on erythrocyte phagocytosis and heme iron recycling within splenic macrophages. To this end, we implemented a dual approach: administering NHS-biotin intravenously to mice, followed by HH exposure, and subsequently introducing PKH67-labelled RBCs into the mice, also followed by HH treatment in vivo. This methodology allowed us to monitor the clearance of RBCs by tracking the retention of both biotin-labelled and PKH67-labelled RBCs using flow cytometry in the blood and spleen of the HH-treated mice. Our findings (as detailed in Figure 4) indicated a pronounced decrease in the presence of both biotin-labelled and PKH67-labelled RBCs in the blood and spleen at 7- and 14-day intervals post HH exposure. Notably, while a natural decline in labelled RBCs over time was expected, the rate of decay was significantly more acute in the NN exposure group as compared to the HH group. This observation was consistent across both blood (Figure 4, A and C; Figure 4, E and G) and spleen samples (Figure 4, B and D; Figure 4, F and H). Furthermore, our analysis revealed that the phagocytic capacity of mouse spleen macrophages toward RBCs was notably diminished following HH exposure, particularly on the 14th day (Figure 4, A-D; Figure 4, E-H). To confirm these findings, we conducted additional assessments of the PKH67-positive RPMs cell population through both flow cytometry and immunofluorescence detection in the spleen post HH exposure. The findings revealed that both the population of splenic RPMs (F4/80hiCD11blo) (Figure 5, A-B) and the PKH67-positive macrophages (Figure 5, A and C; D-E) consistently demonstrated a substantial reduction after 7 or 14 days of HH exposure, further reinforcing the impact of HH on the erythrophagocytic function of splenic macrophages.

Earlier version: HH exposure decreases erythrophagocytosis and iron processing capacity in the spleen.

Revised version: HH exposure decreases RBCs clearance both in the blood and spleen.

Earlier version: Flow cytometry revealed the proportion of FITC-stained cells in the blood (A) and spleen (B) after 7 and 14 days of HH exposure. The proportions of these FITC-positive cells in the blood (C) and spleen (D) are presented as bar graphs. (E) Wright staining of splenic cells after HH exposure for 0 (NN), 7 and 14 days. (F) The proportion of retention RBCs within the spleen cells. (G) Illustrating the experimental design, in which C57BL6 mice (n=3 per group) were treated on days 7 and 11 with a single intravenous Tuftsin injection (1.5 mg/kg), followed by HH exposure until day 14 for spleen isolation and subsequent F4/80 immunohistochemistry and iron staining. (H) F4/80 and iron staining in the spleen after HH exposure for 14 days with Tuftsin treatment. (I) Quantitative the relative fluorescence of F4/80 expression in the spleen described in (H). (J) Semiquantitative iron levels in the spleen described in (H). Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: (A and B) Flow cytometry measurements illustrated the proportions of FITC-stained RBCs in blood and spleen, respectively, after 7 and 14 days of HH exposure. (C and D) Bar graphs depicted the quantified proportions of FITC-positive RBCs in both blood and spleen, respectively. Further, (E and F) presented flow cytometry assessments of the proportions of PKH67-labelled RBCs in blood and spleen following 7 and 14 days of HH exposure. The corresponding proportions of PKH67-positive RBCs in blood (G) and spleen (H) were also shown in bar graph format. The data are expressed as means ±SEM for each experimental group (n = 3 per group). Statistical significance is denoted with * P < 0.05, ** P < 0.01, *** P < 0.001, relative to the NN group or as specified in the graph legends. This compilation of data underlines the impact of HH exposure on the reduction of RBCs clearance in both blood and splenic compartments.

Added in the revised version: HH exposure reduces erythrophagocytosis in the spleen.

Added in the revised version: Mice were administered PKH67-labeled RBCs, followed by exposure to HH for durations of 7 and 14 days. (A) Flow cytometry was employed to analyze the proportions of RPMs (identified as F4/80+CD11b-) and PKH67-positive RPMs in the spleen post-exposure. (B and C) Bar graphs represent the quantified proportions of RPMs (F4/80+CD11b-) and PKH67-positive RPMs in the spleen, respectively. (D) The spleens were subjected to immunofluorescence analysis to detect F4/80 in conjunction with PKH67 fluorescence post-perfusion at 7 and 14 days following HH exposure. (E) Subsequent quantification of the number of PKH67-positive F4/80hi cells in the spleen, as depicted in (D), was conducted. The data derived from these analyses are expressed as means ±SEM for each group (n = 3 per group). Statistical significance is denoted with * P < 0.05, and *** P < 0.001, relative to the indicated group in the graph legends.

Added in the revised version: Reduced erythrophagocytosis leads to RBCs retention in the spleen under HH exposure

Added in the revised version: Based on the reduced RPMs and erythrophagocytosis caused by HH exposure in spleen, we next investigated whether RBCs were retained in spleen after HH exposure. Initial examination involved detecting RBCs in the spleen post HH exposure, both with and without perfusion. HE staining (Figure 6, A-B) and Band 3 immunostaining in situ (Figure 6, C-D, G-H) revealed a significant increase in RBCs numbers in the spleen following 7 and 14 days of HH exposure. Furthermore, we quantified RBCs retention by employing Wright-Giemsa composite staining on single splenic cells post-perfusion at both 7- and 14-days post HH exposure. The results consistently indicated a substantial elevation in RBC counts within the spleen (Figure 6, E-F). To enhance the specificity of our investigation, we labelled RBCs in vitro with PKH67 and then administered them to mice. Following HH exposure for 7 and 14 days, spleen sections were analyzed without perfusion to detect retained PKH67-labeled RBCs. Fluorescence detection techniques was utilized, revealing a marked increase in PKH67-labeled RBCs post HH exposure (Figure 6, I-J). These results collectively confirm a pronounced impairment in erythrophagocytosis within the spleen under HH conditions. This is evidenced by the observed increase in RBCs deformation and retention following 7 and 14 days of HH exposure. The data thus suggest that HH exposure leads to an increase in RBCs retention in the spleen, highlighting significant alterations in splenic function and RBCs dynamics under hypoxic conditions.

Added in the revised version: HH exposure increases RBCs retention in the spleen.

Added in the revised version: (A) HE staining was utilized to examine the presence of RBCs in the spleen post-perfusion at 7 and 14 days following HH exposure. (B) The number of RBCs present in the spleen, as observed in (A), was quantitatively assessed. (C and G) Immunohistochemical and immunofluorescent analyses using Band 3 staining were conducted to evaluate RBCs content in the spleen post-perfusion at 7 and 14 days of HH exposure. (D and H) The expression levels of Band 3 in the spleen, as shown in (C and G), were quantified. (E) Wright-Giemsa composite staining was performed on splenic cells following HH exposure for durations of 0 (NN), 7, and 14 days. (F) The proportion of RBCs within the perfused splenic cell population was determined. (I) Detection of PKH67 fluorescence in the spleen, without perfusion, was conducted after administering PKH67-labeled RBCs and subjecting them to HH exposure for 7 and 14 days. (J) The number of PKH67-positive RBCs within the spleen, as outlined in (I), was quantified. Data presented in this figure are expressed as means ±SEM for each experimental group (n = 3 per group). Statistical significance is indicated by * P < 0.05, ** P < 0.01, *** P < 0.001 when compared to the NN group or as specified.

Added in the revised version: HH induced reduction in erythrophagocytosis leads to a decrease in the capacity of iron processing in the spleen

Added in the revised version: To investigate the hypothesis that increased retention of RBCs in the spleen due to HH exposure is a result of impaired erythrophagocytosis by RPMs, we employed a series of experiments. Initially, RBCs were labelled with PKH67 cell linker and then injected into mice. Subsequently, Tuftsin was administered to stimulate the phagocytic activity of macrophages. After 1 h, 1 day, 3 days, 5 days, or 7 days of NN exposure, PKH67 fluorescence was analyzed via flow cytometry. The findings demonstrated a gradual reduction in PKH67 fluorescence over the course of NN exposure; however, a significant decrease in PKH67 fluorescence was noted in the spleen following Tuftsin administration compared to the NN group. This observation suggested that RBC lifespan is diminished following enhanced erythrophagocytosis (as presented in Figure 7, A-B). Further investigation was conducted on the ability of the spleen to regulate heme iron recycling under HH conditions through immunofluorescence and iron staining (Figure 7C). The results, as illustrated in Figure 7D, indicated that both F4/80 expression and iron deposition in the red pulp of the spleen were notably reduced following HH exposure, compared to the NN group. Contrastingly, Tuftsin administration under HH conditions resulted in an upsurge in F4/80 expression and iron deposition in the red pulp (Figure 7, E-F). These findings collectively suggest that the splenic capacity for erythrocyte phagocytosis and subsequent heme iron recycling is significantly compromised under HH conditions, primarily due to a reduction in RPMs. This reduced erythrophagocytic capacity of the spleen under HH exposure is a critical factor contributing to the decreased efficiency of iron processing and increased RBCs retention in this organ.

Added in the revised version: Impaired erythrophagocytosis of RBCs after HH exposure leads to a decrease in iron processing capacity in the spleen.

Added in the revised version: To stimulate phagocytosis of macrophages, Tuftsin was administered immediately following the injection of PKH67-labeled RBCs into mice. Subsequent to various exposure durations under NN conditions (1 h, 1, 3, 5, or 7 days), (A) PKH67 fluorescence within the spleen was measured using flow cytometry. (B) The percentages of PKH67-positive RBCs in the spleen are represented in bar graph format. (C) The experimental design is depicted, wherein C57BL/6 mice (n = 3 per group) were administered a single intravenous dose of Tuftsin (0.15 mg/kg) on days 7 and 11, followed by HH exposure until day 14, culminating in spleen isolation for F4/80 immunohistochemistry and iron staining. (D) Demonstrates F4/80 and iron staining in the spleen after 14 days of HH exposure with Tuftsin treatment. (E) Quantitative analysis of the relative fluorescence intensity of F4/80 expression in the spleen as described in (D). (F) Semi-quantitative assessment of iron levels in the spleen as outlined in (D). Data presented in this figure are articulated as means ± SEM for each group (n = 3 per group). Statistical significance is denoted by * P < 0.05, ** P < 0.01, *** P < 0.001 when compared to the NN group or as indicated.

HH exposure induces iron mobilization and ferroptosis in the spleen

Removed in the revised version: To elucidate the precise mechanisms underlying the macrophage reduction caused by HH, we examined the expression of proteins related to iron metabolism and ferroptosis in mice treated with HH and analyzed the corresponding gene expression in peripheral blood mononuclear cells (PBMCs) from healthy humans acutely exposed to HA conditions using GEO data (No. GSE46480). The results showed that, compared to the NN group, the expression levels of HO-1, Ft-L, Ft-H, NCOA4, and xCT were decreased, while Fpn, TfR, and ACSL4 expressions were significantly increased after 7 and 14 days of HH exposure (Figure 5, A-B, D-E; Figure S1, A-D).

Removed in the revised version: HH exposure promotes iron mobilization and induces ferroptosis in the spleen.

Removed in the revised version: The C57BL/6 mice were treated with NN and HH for 7 days, and the spleen was collected for subsequent detection. (A) Western blot detection of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression in spleen. (B) Statistical analysis of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression. (C) GEO data analysis of HMOX1, FTL, FTH1, NCOA4, SLC40A1, and TFRC mRNA expression in PMBCs before and after climbing to high altitude (n=98). Western blot detection of ACSL4, GPX4, and xCT protein expression in the spleen is shown in (D) and its statistical analysis in (E). (F) GEO data analysis of ACSL4, GPX4 and SLC7A11 mRNA expression in PMBCs before and after reaching high altitude (n=98). (G) The content of Fe2+ in the spleen was detected using the FerroFarRed probe by flow cytometry. (H) The level of lipid ROS in the spleen was detected using the C11-BODIPY probe by flow cytometry. The MDA (I), Cys (J) and GSH (K) levels in the spleen were detected using biochemical detection kits, respectively. (L) Lillie staining of Fe2+ in the spleen after 7 days of HH exposure. (M) Fe2+ deposition in the spleen as described in (L) was quantified. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Earlier version: With the exception of NCOA4, the alterations in gene expressions related to iron metabolism and ferroptosis in PBMCs were consistent with our Western blot results (Figure 5, C and F). The GPX4 gene and protein expressions remained largely unchanged in both human PBMCs and mouse spleens. The changes in iron metabolism-and ferroptosis-related genes in PBMCs reflect the protein changes in the spleen under HH exposure. We detected Fe2+ and lipid ROS levels in the spleen by flow cytometry, and quantified MDA, GSH, and Cys levels using biochemical detection kits. As depicted in Figure 5, G and H; Figure S1, I-K, the Fe2+ (Figure 5G; Figure S1, E and F) and lipid ROS (Figure 5H; Figure S1, G and H) levels in the spleen increased significantly after HH exposure. Additionally, the MDA content (Figure 5I; Figure S1I) in the spleen increased, whereas Cys (Figure 5J; Figure S1J) and GSH (Figure 5K; Figure S1K) levels significantly decreased after HH exposure. To ascertain the increased Fe2+ primarily emanated from the red pulp, we detected Fe2+ deposition and distribution in the spleen by Lillie stain following 7 and 14 days of HH exposure (Figure 5, L-M; Figure S1, L-M). These results demonstrated increased Fe2+ primarily deposited in the red pulp of the spleen. Taken together, these findings suggest that iron mobilization in the spleen was enhanced and ferroptosis was induced after 7 days of HH exposure.

Revised version: To elucidate the precise mechanisms underlying the macrophage reduction caused by HH, we examined the expression of proteins related to iron metabolism and ferroptosis in mice treated with HH and analyzed the corresponding gene expression in peripheral blood mononuclear cells (PBMCs) from healthy humans acutely exposed to HA conditions using GEO data (No. GSE46480). The results showed that, compared to the NN group, the expression levels of HO-1, Ft-L, Ft-H, NCOA4, and xCT were decreased, while Fpn, TfR, and ACSL4 expressions were significantly increased after 7 and 14 days of HH exposure (Figure 8, A-B, D-E; Figure S3, A-D). Apart from NCOA4, the alterations in gene expressions related to iron metabolism and ferroptosis in PBMCs were consistent with our Western blot results (Figure 8, C and F). The GPX4 gene and protein expressions remained largely unchanged in both human PBMCs and mouse spleens. The changes in iron metabolism- and ferroptosis-related genes in PBMCs reflect the protein changes in the spleen under HH exposure. We detected Fe2+ and lipid ROS levels in the spleen by flow cytometry, and quantified MDA, GSH, and Cys levels using biochemical detection kits. As depicted in Figure 8, G and H; Figure S3, I-K, the Fe2+ (Figure 8G; Figure S3, E and F) and lipid ROS (Figure 8H; Figure S3, G and H) levels in the spleen increased significantly after HH exposure. Additionally, the MDA content (Figure 8I; Figure S3I) in the spleen increased, whereas Cys (Figure 8J; Figure S3J) and GSH (Figure 8K; Figure S3K) levels significantly decreased after HH exposure. To ascertain the increased Fe2+ primarily emanated from the red pulp, we detected Fe2+ deposition and distribution in the spleen by Lillie stain following 7 and 14 days of HH exposure (Figure 8, L-M; Figure S3, L-M). These results demonstrated increased Fe2+ primarily deposited in the red pulp of the spleen. Taken together, these findings suggest that iron mobilization in the spleen was enhanced and ferroptosis was induced after 7 days of HH exposure.

Added in the revised version: HH exposure enhances iron mobilization and induces ferroptosis in the spleen.

Added in the revised version: The C57BL/6 mice were treated with NN and HH for 7 days, and the spleen was collected for subsequent detection. (A) Western blot detection of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression in spleen. (B) Statistical analysis of HO-1, Ft-L, Ft-H, NCOA4, Fpn and TfR protein expression. (C) GEO data analysis of HMOX1, FTL, FTH1, NCOA4, SLC40A1, and TFRC mRNA expression in PMBCs before and after climbing to HA (n=98). (D) Western blot analysis for ACSL4, GPX4, and xCT protein expression in the spleen, with (E) depicting the statistical analysis of these protein levels. (F) GEO data analysis of ACSL4, GPX4 and SLC7A11 mRNA expression in PMBCs before and after reaching HA (n=98). (G) The content of Fe2+ in the spleen was detected using the FerroFarRed probe by flow cytometry. (H) The level of lipid ROS in the spleen was detected using the C11-BODIPY probe by flow cytometry. The MDA (I), Cys (J) and GSH (K) levels in the spleen were detected using biochemical detection kits, respectively. (L) Lillie staining of Fe2+ in the spleen after 7 days of HH exposure. (M) Fe2+ deposition in the spleen as described in (L) was quantified. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Hypoxia induces ferroptosis in primary splenic macrophages

Earlier version: To investigate whether hypoxia exposure can trigger ferroptosis in macrophages, we measured alterations in Fe2+ and lipid ROS levels, cell viability, phagocytosis, and ferroptosis-related protein expressions in primary splenic macrophages under hypoxia in the presence of a ferroptosis inhibitor (Fer-1). As depicted in Figure 6, compared with the normoxia control group (Nor), Fe2+ (Figure 6, A-C) and lipid ROS (Figure 6, D-E) levels significantly increased, whereas the number of viable cells (Figure 6G-H) and phagocytosis ability (Figure 6I) significantly decreased after 24 h of hypoxia (Hyp) treatment. In addition, prolonged hypoxia exposure led to a gradual increase in ACSL4 expression, while xCT and GPX4 expressions decreased (Figure 6F). Fer-1 mitigated the increase in Fe2+ content (Figure 6, J-K) and reversed the expression changes in ACSL4, xCT, and GPX4 induced by hypoxia exposure (Figure 6, L-O). Simultaneously, Fer-1 reversed the alterations in MDA, Cys, and GSH levels induced by hypoxia (Figure 6, P-R). These findings confirm that hypoxia induced ferroptosis in primary splenic macrophages.

Revised version: To investigate whether hypoxia exposure can trigger ferroptosis in macrophages, we measured alterations in Fe2+ and lipid ROS levels, cell viability, phagocytosis, and ferroptosis-related protein expressions in primary splenic macrophages under hypoxia in the presence of a ferroptosis inhibitor (Fer-1). As depicted in Figure 9, compared with the normoxia control group (Nor), Fe2+ (Figure 9, A-C) and lipid ROS (Figure 9, D-E) levels significantly increased, whereas the number of viable cells (Figure 9, G-H) and phagocytosis ability (Figure 9I) significantly decreased after 24 h of hypoxia (Hyp) treatment. In addition, prolonged hypoxia exposure led to a gradual increase in ACSL4 expression, while xCT and GPX4 expressions decreased (Figure 9F). Fer-1 mitigated the increase in Fe2+ content (Figure 9, J-K) and reversed the expression changes in ACSL4, xCT, and GPX4 induced by hypoxia exposure (Figure 9, L-O). Simultaneously, Fer-1 reversed the alterations in MDA, Cys, and GSH levels induced by hypoxia (Figure 9, P-R). These findings confirm that hypoxia induced ferroptosis in primary splenic macrophages.

Earlier version: Hypoxia enhances primary peritoneal macrophage ferroptosis.

Revised version: Hypoxia enhances the ferroptosis in splenic macrophage.

Earlier version: Primary peripheral macrophages were cultured under 1% hypoxia for varying durations (0, 12, 24, 48 h), either with or without a pre-treatment of 10 µM Fer-1 for 1 h, before being collected for further examination. (A) Immunofluorescence detection of Fe2+ levels in macrophages after hypoxia exposure for 24 h using the FerroFarRed probe under a fluorescence microscope. (B) Quantitative results of fluorescence intensity in the macrophages described in A. (C) Flow cytometry detection of Fe2+ levels in macrophages exposed to hypoxia for 24 h using the FerroFarRed probe. (D) Immunofluorescence detection of lipid ROS levels in macrophages exposed to hypoxia for 24 h using the C11-BODIPY probe under a fluorescence microscope. (E) Quantitative results of mean fluorescence intensity in the macrophages described in D. (F) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages under hypoxia for different times. (G) Flow cytometry was employed to analyze Calcein/PI double staining in macrophages following a 24 h hypoxia exposure. (H) Cell death proportions in macrophages are given as bar graphs. (I) Macrophage phagocytic activity against Cy5.5-labelled E.coli following 24 h of hypoxia exposure was assessed using flow cytometry in vitro. (J) Fe2+ levels in macrophages pre-treated with Fer-1, then exposed to 24 h of hypoxia, were detected via the FerroFarRed probe in immunofluorescence. (K) Quantitative results of fluorescence intensity in the macrophages described in J. (L) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages pre-treated with Fer-1, then followed by 24 h of hypoxia exposure. (M-O) Statistical analysis of ACSL4, xCT, and GPX4 protein expression in L. MDA (P), Cys (Q), and GSH (R) levels in macrophages pre-treated with Fer-1 and subsequently exposed to 24 h of hypoxia were detected using kits and biochemical methods. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Revised version: Splenic macrophages were cultured under 1% hypoxia for varying durations (0, 12, 24, 48 h), either with or without a pre-treatment of 10 µM Fer-1 for 1 h, before being collected for further examination. (A) Immunofluorescence detection of Fe2+ levels in macrophages after hypoxia exposure for 24 h using the FerroFarRed probe under a fluorescence microscope. (B) Quantitative results of fluorescence intensity in the macrophages described in A. (C) Flow cytometry detection of Fe2+ levels in macrophages exposed to hypoxia for 24 h using the FerroFarRed probe. (D) Immunofluorescence detection of lipid ROS levels in macrophages exposed to hypoxia for 24 h using the C11-BODIPY probe under a fluorescence microscope. (E) Quantitative results of mean fluorescence intensity in the macrophages described in (D). (F) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages under hypoxia for different times. (G) Flow cytometry was employed to analyze Calcein/PI double staining in macrophages following a 24 h hypoxia exposure. (H) Cell death proportions in macrophages are given as bar graphs. (I) Macrophage phagocytic activity against Cy5.5-labelled E.coli following 24 h of hypoxia exposure was assessed using flow cytometry in vitro. (J) Fe2+ levels in macrophages pre-treated with Fer-1, then exposed to 24 h of hypoxia, were detected via the FerroFarRed probe in immunofluorescence. (K) Quantitative results of fluorescence intensity in the macrophages described in (J). (L) Western blot detection of ACSL4, GPX4 and xCT protein expression in macrophages pre-treated with Fer-1, then followed by 24 h of hypoxia exposure. (M-O) Statistical analysis of ACSL4, xCT, and GPX4 protein expression in L. MDA (P), Cys (Q), and GSH (R) levels in macrophages pre-treated with Fer-1 and subsequently exposed to 24 h of hypoxia were detected using kits and biochemical methods. Data are expressed as the means ± SEM (n = 3 per group); * P < 0.05, ** P < 0.01, *** P < 0.001 versus the NN group or the indicated group.

Discussion

Earlier version: The purpose of this study was to explore the effects and mechanisms of HA/HH exposure on erythrophagocytosis and iron circulation in mouse spleens. Here, we used an HH chamber to simulate 6000 m HA exposure and found that HH exposure induced iron mobilization and activated ferroptosis in the spleen, especially in red pulp macrophages (RPMs), which subsequently inhibited the phagocytosis and clearance of RBCs. Finally, chronic exposure to HA/HH may promote the retention of RBCs in the spleen, cause splenomegaly, advance RBC production, and promote the occurrence and development of HAPC (Figure 7).

Revised version: The purpose of this study was to explore the effects and mechanisms of HA/HH exposure on erythrophagocytosis and iron circulation in mouse spleens. Here, we used an HH chamber to simulate 6000 m HA exposure and found that HH exposure induced iron mobilization and activated ferroptosis in the spleen, especially in RPMs, which subsequently inhibited the phagocytosis and clearance of RBCs. Finally, chronic exposure to HA/HH may promote the retention of RBCs in the spleen, cause splenomegaly, advance RBC production, and promote the occurrence and development of HAPC.

Removed in the revised version: High altitude/hypobaric hypoxia exposure suppresses erythrophagocytosis by inducing macrophages ferroptosis in the spleen.

Removed in the revised version: HA/HH exposure induces iron mobilization, especially increased Fe2+ levels, which are further increased in the spleen/macrophages of mice. On the other hand, HA/HH exposure increases ACSL4 expression, suppressing xCT expression and GSH production, which escalates lipid peroxidation and subsequent macrophage ferroptosis. Meanwhile, HA/HH exposure decreases monocyte migration and differentiation from the BM to the spleen. Taken together, HA/HH exposure inhibits erythrophagocytosis and further promotes erythrocytosis, which may aggravate HAPC development.

Earlier version: RBCs/erythrocytes are the carriers of oxygen and the most abundant cell type in the body. Erythrocytes are rapidly increased by triggering splenic contraction in acute HA/HH exposure [21] and stimulating erythropoiesis in subsequent continuous or chronic HA/HH exposure. Changes in spleen morphology and size are closely related to the spleen’s ability to recover RBCs [22], and studies have shown that the spleen is in a contraction state after short-term exposure to HA/HH, and approximately 40% of the increase in RBCs is due to the contraction of the spleen [21, 23, 24]. In the present study, mice were treated under HH to mimic HA exposure (mainly mimicking the low oxygen partial pressure environment of 6000 m HA) for different times according to other studies [25, 26]. Our results showed that HH exposure not only significantly increased the RBC and HGB contents but also significantly increased CD47+ cells in blood. However, short-term (1 d) HH exposure caused spleen contraction and induced splenomegaly after 3 d of HH treatment. This contraction can trigger an immediate release of RBCs into the bloodstream in instances of substantial blood loss or significant reduction of RBCs. Moreover, elevated oxygen consumption rates in certain animal species can be partially attributed to splenic contractions, which augment hematocrit levels and the overall volume of circulating blood, thereby enhancing venous return and oxygen delivery [27, 28]. Thus, we hypothesized that the body, under such conditions, is incapable of generating sufficient RBCs promptly enough to facilitate enhanced oxygen delivery. Consequently, the spleen reacts by releasing its stored RBCs through splenic constriction, leading to a measurable reduction in spleen size. These results are consistent with other research results; that is, HH or HA exposure affects spleen morphology and further 483 affects RBC counts and HGB levels [29–31]. Considering these findings, the present study postulates that hypoxia-induced inhibition of erythrophagocytosis may lead to RBC retention. However, we acknowledge that the manuscript in its current preprint form does not offer conclusive evidence to substantiate this hypothesis. To bridge this gap, we further conducted experiments where the spleen was perfused, and total cells were collected post HH exposure. These cells were then smeared onto slides and subjected to Wright staining. Our results unequivocally demonstrate an evident increase in deformation and retention of RBCs in the spleen following 7 and 14 days of HH exposure. This finding strengthens our initial hypothesis and contributes a novel perspective to the understanding of splenic responses under hypoxic conditions.

Revised version: RBCs/erythrocytes are the carriers of oxygen and the most abundant cell type in the body. Erythrocytes are rapidly increased by triggering splenic contraction in acute HA/HH exposure [23] and stimulating erythropoiesis in subsequent continuous or chronic HA/HH exposure. Changes in spleen morphology and size are closely related to the spleen’s ability to recover RBCs [24], and studies have shown that the spleen is in a contraction state after short-term exposure to HA/HH, and approximately 40% of the increase in RBCs is due to the contraction of the spleen [23, 25, 26]. In the present study, mice were treated under HH to mimic HA exposure (mainly mimicking the low oxygen partial pressure environment of 6000 m HA) for different times according to other studies [27, 28]. Our results showed that HH exposure significantly increased the number of RBCs and HGB contents. However, short-term (1 day) HH exposure caused spleen contraction and induced splenomegaly after 3 days of HH treatment. This contraction can trigger an immediate release of RBCs into the bloodstream in instances of substantial blood loss or significant reduction of RBCs. Moreover, elevated oxygen consumption rates in certain animal species can be partially attributed to splenic contractions, which augment haematocrit levels and the overall volume of circulating blood, thereby enhancing venous return and oxygen delivery [29, 30]. Thus, we hypothesized that the body, under such conditions, is incapable of generating sufficient RBCs promptly enough to facilitate enhanced oxygen delivery. Consequently, the spleen reacts by releasing its stored RBCs through splenic constriction, leading to a measurable reduction in spleen size. These results are consistent with other research results; that is, HH or HA exposure affects spleen morphology and further affects RBCs counts and HGB levels [31–33].

Earlier version: HAPC is a common chronic HA disease characterized by excessive proliferation of RBCs caused by HH conditions [32]. Under physiological conditions, RBC homeostasis is maintained by balancing erythropoiesis and erythrocyte clearance [4]. The ability of RBC disposal in the spleen for iron recycling under HA/HH exposure was important to RBC regeneration in bone marrow (BM) [33]. Thus, we hypothesized that erythrophagocytosis and iron recycling in the spleen were altered by HA/HH exposure and further disturbed RBC homeostasis and affected the progression of HAPC. We found that compared with sham group mice, HH significantly increased the contents of RBCs and HGB in the blood of splenectomy mice. This strongly verified that the spleen is a key organism that maintains RBC magnitude within a certain range of physiological statuses under HA/HH conditions.

Revised version: HAPC is a common chronic HA disease characterized by excessive proliferation of RBCs caused by HH conditions [34]. Under physiological conditions, RBC homeostasis is maintained by balancing erythropoiesis and erythrocyte clearance [4]. The ability of RBC disposal in the spleen for iron recycling under HA/HH exposure was important to RBC regeneration in BM [35]. Thus, we hypothesized that erythrophagocytosis and iron recycling in the spleen were altered by HA/HH exposure and further disturbed RBC homeostasis and affected the progression of HAPC. We found that compared with sham group mice, HH significantly increased the contents of RBCs and HGB in the blood of splenectomy mice. This strongly verified that the spleen is a key organism that maintains RBC magnitude within a certain range of physiological statuses under HA/HH conditions.

Earlier version: As originally proposed by Metchnikoff in the 19th century, macrophages, especially RPMs, play a pivotal role in the regulation of RBCs and iron homeostasis [34, 35]. We detected the macrophage population in the spleen and found that HH exposure for 7 and 14 days reduced the total macrophage and M1-and M2-type macrophage proportions in the spleen. This result is supported by another study, that is, hypoxia reduces M1-type macrophages in human gastric cancer [36]. We next observed whether monocyte migration and differentiation from the BM to the spleen replenish reduced macrophages and sustain macrophage homeostasis after HH treatment. It is well known that splenic erythropoiesis is a dynamic process [4]. Upon phagocytosing erythrocytes in macrophages, spleen tissue (including macrophages and fibroblasts) produces the chemokine C–C motif chemokine ligand 2 and 7 (CCL2 and CCL7) to recruit blood monocytes to the spleen [37, 38]. Unlike most studies in tumours, which mainly suggested that macrophage accumulation in hypoxic areas was derived by enhanced chemokine expression in tumours [39, 40], CCL2, CCL7 and Csf1 expression was decreased in the spleen after HH exposure for 7 and 14 days in our study. Furthermore, we found that HH exposure inhibited Ly6Chi monocyte migration from the BM to spleen, where these cells subsequently differentiated into RPMs. The combination of decreased splenic RPMs, along with the reduced BM-derived Ly6Chi monocytes, suggests that the homeostasis of the RPM population of circulating monocytes was also inhibited after the RPM depletion induced by HH exposure. The observed decrease in monocytes within the BM is likely attributable to the fact that monocytes and precursor cells for RBCs both originate from the same hematopoietic stem cells within the BM [41]. As such, the differentiation to monocyte is reduced under hypoxic conditions, which may subsequently cause a decrease in migration to spleen. Furthermore, we hypothesize that an increased migration of monocytes to other tissues under HH exposure may also contribute to the decreased migration to the spleen. The liver, which partially contributes to the clearance of RBCs, may play a role in this process. Our investigations to date have indeed identified an increased monocyte migration to the liver. We were pleased to discover an elevation in CSF1 expression in the liver following HH exposure for both 7 and 14 days. This finding was corroborated through flow cytometry, which confirmed an increase in monocyte migration to the liver.

Revised version: As originally proposed by Metchnikoff in the 19th century, macrophages, especially RPMs, play a pivotal role in the regulation of RBCs and iron homeostasis [36, 37]. We detected the macrophage population in the spleen and found that HH exposure for 7 and 14 days reduced the total macrophages in the spleen. We next observed whether the migration and differentiation of monocytes from the BM to the spleen supplemented the reduced macrophages after HH treatment and maintained their homeostasis. It is well known that splenic erythrophagocytosis is a dynamic process [4]. Upon phagocytosing erythrocytes in macrophages, spleen tissue (including macrophages and fibroblasts) produces the chemokine C–C motif chemokine ligand 2 and 7 (CCL2 and CCL7) to recruit blood monocytes to the spleen [38, 39]. In our study, the expression of CCL2, CCL7, and Csf1 in the spleen decreased after 7 and 14 days of HH exposure. Furthermore, we found that HH exposure inhibited Ly6Chi monocyte migration from the BM to spleen, where these cells subsequently differentiated into RPMs. The combination of decreased splenic RPMs, along with the reduced BM-derived Ly6Chi monocytes, suggests that the homeostasis of the RPMs population in circulating monocytes was also inhibited after the depletion of RPMs induced by HH exposure. The observed decrease in monocytes within the BM is likely attributable to the fact that monocytes and precursor cells for RBCs both originate from the same hematopoietic stem cells within the BM [40]. As such, the differentiation to monocyte is reduced under hypoxic conditions, which may subsequently cause a decrease in migration to spleen.

Added in the revised version: HO-1 is a pivotal antioxidant and cytoprotective enzyme, instrumental in catalyzing the degradation of heme, thereby maintaining heme homeostasis and mitigating free heme-induced toxicity [41]. Notably, it has been posited that mice deficient in HO-1 exhibit a significant depletion of RPMs, emphasizing the important role of this enzyme in the development and maintenance of RPM after RBCs clearance [42]. In our study, the observed decrease in HO-1 protein levels in the red pulp suggests a potential reduction in the number of macrophages induced by HH exposure. This inference is made despite the general understanding that HO-1 expression is typically upregulated by hypoxia-inducible factor 1 (HIF-1) under hypoxic conditions [43, 44]. Along with the decrease in the number of macrophages caused by HH, we also noticed that the phagocytic ability of macrophages in the spleen is inhibited, as evidenced both in vitro and in vivo. Furthermore, the administration of tufftsin markedly enhanced heme iron processing in the spleen under HH conditions. However, our findings appear to diverge from those reported by Anand et al., which suggested an increase in macrophage phagocytosis under hypoxic conditions, mediated by HIF-1α [45]. This discrepancy may stem from the differing methodologies employed; whereas Anand et al. utilized intermittent hypoxia in their research, our study was characterized by consistent hypoxia. However, it should be acknowledged that using HO-1 expression as an alternative marker to quantify macrophage numbers is not without limitations. Variations in HO-1 expression per cell can yield misleading results, and the localization of this effect to the red pulp does not definitively equate to a conclusion applicable to macrophages, given the inherent heterogeneity of this region and the spleen overall. Considering the multifaceted nature of spleen function and cellular dynamics under hypoxic stress, this complexity requires careful interpretation of our data.

Earlier version: HO-1 is a major antioxidant and cytoprotective enzyme that catalyzes the degradation of heme to provide heme homeostasis and protect against free heme-induced toxicity [42]. In addition, it was supposed that mice deficient in HO-1 are largely devoid of RPMs and suggested that HO-1 is critical for RPM development and survival after erythrocyte clearance [43]. The reduced HO-1 protein of red pulp in our study also implied decreased macrophage accounts induced by HH exposure, although HO-1 was considered to be upregulated by hypoxia-inducible factor 1 (HIF-1) under hypoxia [44, 45]. Along with the decreased macrophage amount induced by HH, the phagocytic ability of macrophages in the spleen was also inhibited both in vitro and in vivo. Meanwhile, the injection of tufftsin strongly confirmed the enhancement of heme iron processing in the spleen under HH exposure. Nevertheless, our data conflict with the study of Anand et al, which reported that hypoxia causes an increase in phagocytosis by macrophages in a HIF-1α-dependent manner [46]. This may be because they performed intermittent hypoxia, while we used consistent hypoxia in the present study. Macrophages play a crucial role in maintaining erythrocyte homeostasis, partially due to their function in digesting the HGB content of cleared RBCs and recycling the iron back to erythroid progenitors for heme synthesis and HGB production [47]. In mammals, the majority of bodily iron exists in the form of heme, with the most substantial pool comprising HGB [48]. HGB contains four prosthetic hemoglobins, and each mature RBC is thought to contain approximately 1.2 × 109 heme moieties [47]. As previously mentioned, the majority of iron required to sustain erythropoiesis is derived from recycled RBCs. In general, RPMs are equipped with molecular machinery capable of neutralizing the toxic effects of heme and metabolizing iron [49, 50]. Nonetheless, any disruptions in the process of erythrophagocytosis, the uptake and degradation of erythrocytes by macrophages, could potentially result in abnormal iron metabolism, potentially manifesting as conditions such as anemia or iron overload [49]. The release of heme upon processed RBCs constitutes a permanent and considerable threat for iron cytotoxicity in macrophages and may eventually result in a specific form of programmed cell death termed ferroptosis [51], which may cause decreased cell counts [52, 53].

Revised version: Macrophages play a crucial role in maintaining erythrocyte homeostasis, partially due to their function in digesting the HGB content of cleared RBCs and recycling the iron back to erythroid progenitors for heme synthesis and HGB production [46]. In mammals, most of the bodily iron exists in the form of heme, with the most substantial pool comprising HGB [47]. HGB contains four prosthetic hemoglobins, and each mature RBC is thought to contain approximately 1.2 × 109 heme moieties [46]. As previously mentioned, most of the iron required to sustain erythropoiesis is derived from recycled RBCs. In general, RPMs are equipped with molecular machinery capable of neutralizing the toxic effects of heme and metabolizing iron [48, 49]. Nonetheless, any disruptions in the process of erythrophagocytosis, the uptake and degradation of erythrocytes by macrophages, could potentially result in abnormal iron metabolism, potentially manifesting as conditions such as anemia or iron overload [48]. The release of heme upon processed RBCs constitutes a permanent and considerable threat for iron cytotoxicity in macrophages and may eventually result in a specific form of programmed cell death termed ferroptosis [50], which may cause decreased cell counts [51, 52].

Earlier version: Accordingly, we investigated the iron metabolism status and explored ferroptosis in spleen/macrophages after HH/hypoxia treatment and found that HH exposure prompted iron mobilization, especially in ferritinophagy and lipid peroxidation in the spleen. Interestingly, we found that except for NCOA4, all the gene expression changes in the PBMCs of humans after acute HA exposure were similar to the changes in the spleens of mice after HH exposure for 7 and 14 days. This is probably because NCOA4 in the spleen mediates iron release from Ft-L and facilitates iron reuse, while PBMCs do not possess this function in vivo. These results not only indicated that HH exposure induces spleen ferroptosis but also implied that the gene expression changes in PBMCs under HA/HH may reflect RBC processing functions in the spleen and further indicate the iron metabolism status in the clinic. We further found that the exact mechanism of macrophage ferroptosis induced by HH exposure was caused by increased Fe2+ and decreased antioxidative system expression, which finally resulted in lipid peroxidation of macrophages. It has been proposed that heme catabolism by HO-1 protects macrophages from oxidative stress [49]. The enhanced ROS and lipid peroxidation in vivo were also consistent with the decreased HO-1 expression in the spleen. In addition, 1% hypoxia treatment induced ferroptosis in vitro, which was reduced by treatment with ferrostatin-1, a ferroptosis inhibitor. However, our data were inconsistent with the study of Fuhrmann et al., which reported that hypoxia inhibits ferritinophagy and protects against ferroptosis [54]. We hypothesized that the enhanced iron demand for new RBC regeneration in BM caused by HH leads to an increase in NCOA4 and Ft-L protein expression in the spleen. On the other hand, hypoxia inhibited anti-ferroptosis system expression, including reduced Gpx4 expression and increased lipid ROS production. The results were supported by the group of Youssef LA et al, who reported that increased erythrophagocytosis induces ferroptosis in RPMs in a mouse model of transfusion [55]. However, as most studies have reported, ferroptosis is not only characterized by increased lipid ROS but also significantly shrinking mitochondria [56]. We never found shrinking mitochondria in RPMs after HH exposure (data not shown). Gao et al. proposed that mitochondria played a crucial role in cysteine deprivation-induced ferroptosis but not in that induced by inhibiting glutathione peroxidase-4 (GPX4), the most downstream component of the ferroptosis pathway [57]. Interestingly, in our in vivo study, Gpx4 expression was also not changed after HH exposure. It is still not clear whether mitochondria are involved in ferroptosis. Whether HH treatment caused mitochondrial swelling directly and the exact mechanism involved in this process still need to be further investigated.

Revised version: Accordingly, we investigated the iron metabolism status and explored ferroptosis in spleen/macrophages after HH/hypoxia treatment and found that HH exposure prompted iron mobilization, especially in ferritinophagy and lipid peroxidation in the spleen. Interestingly, we found that all gene expression changes in PBMCs after acute HA exposure in humans, except for NCOA4, were similar to those in the spleen of mice exposed to HH for 7 and 14 days. This is probably because NCOA4 in the spleen mediates iron release from Ft-L and facilitates iron reuse, while PBMCs do not possess this function in vivo. These results not only indicated that HH exposure induces spleen ferroptosis but also implied that the gene expression changes in PBMCs under HA/HH may reflect RBC processing functions in the spleen and further indicate the iron metabolism status in the clinic. We further found that the exact mechanism of macrophage ferroptosis induced by HH exposure was caused by increased Fe2+ and decreased antioxidative system expression, which finally resulted in lipid peroxidation of macrophages. It has been proposed that heme catabolism by HO-1 protects macrophages from oxidative stress [48]. The enhanced ROS and lipid peroxidation in vivo were also consistent with the decreased HO-1 expression in the spleen. In addition, 1% hypoxia treatment induced ferroptosis in vitro, which was reduced by treatment with ferrostatin-1, a ferroptosis inhibitor. However, our data were inconsistent with the study of Fuhrmann et al., which reported that hypoxia inhibits ferritinophagy and protects against ferroptosis [53]. We hypothesized that the enhanced iron demand for new RBCs regeneration in BM caused by HH leads to a decrease in the expression of NCOA4 and Ft-L proteins in the spleen. On the other hand, hypoxia inhibited anti-ferroptosis system expression, including reduced Gpx4 expression and increased lipid ROS production. The results were supported by the group of Youssef LA et al, who reported that increased erythrophagocytosis induces ferroptosis in RPMs in a mouse model of transfusion [54]. However, as most studies have reported, ferroptosis is not only characterized by increased lipid ROS but also significantly shrinking mitochondria [55]. We never found shrinking mitochondria in RPMs after HH exposure (data not shown). Gao et al. proposed that mitochondria played a crucial role in cysteine deprivation-induced ferroptosis but not in that induced by inhibiting glutathione peroxidase-4 (GPX4), the most downstream component of the ferroptosis pathway [56]. Interestingly, in our in vivo study, Gpx4 expression was also not changed after HH exposure. It is still not clear whether mitochondria are involved in ferroptosis. Whether HH treatment caused mitochondrial swelling directly and the exact mechanism involved in this process still need to be further investigated.

Added in the revised version: It is also possible that the spleen may adapt to the increased load of RBCs by provide new avenues for treating HAPC and other conditions related to HA exposure.

Removed in the revised version: When the splenic macrophages’ phagocytic function is compromised, the burden of managing the increased number of RBCs falls to the other cells within the spleen, as well as other clearance mechanisms within the body [58]. The endothelium, for instance, has been shown to be involved in the removal of aged and damaged RBCs [59]. The sinusoidal endothelial cells can recognize these RBCs, which often have altered surface properties, leading to their removal from circulation [59, 60]. Therefore, under HH conditions, endothelial cells may play a compensatory role in RBC clearance when macrophage function is impaired. Other macrophages, such as liver Kupffer cells and bone marrow macrophages, may also contribute to managing this buildup of RBCs [61]. Kupffer cells are the resident macrophages of the liver and have significant phagocytic ability. Similarly, macrophages in the bone marrow play a crucial role in iron recycling from aged erythrocytes [61, 62]. Therefore, these macrophages could potentially compensate for the reduction in splenic erythrophagocytosis, although their ability to effectively handle a large surge in RBC numbers remains to be determined.

Removed in the revised version: It is also possible that the spleen may adapt to the increased load of RBCs by expanding the red pulp, leading to splenomegaly [63]. The red pulp is highly vascular and contains numerous sinuses lined with macrophages [64]. This expansion could potentially accommodate a greater number of RBCs, thus enhancing the spleen’s capacity for erythrophagocytosis and iron recycling. However, these compensatory mechanisms may be insufficient or become overwhelmed over time, leading to pathological conditions such as HAPC. Further research is necessary to fully understand the consequences of splenic ferroptosis on RBC homeostasis and how these effects could be mitigated or reversed. Finally, as these observations are based on experimental models, they should be corroborated in human studies to assess their clinical relevance and potential implications for managing conditions related to high altitude exposure. The potential for therapeutic interventions that target ferroptosis, erythrophagocytosis, and iron recycling should also be explored, as these could provide new avenues for treating HAPC and other conditions related to high altitude exposure.

Earlier version: Based on the above experiments, we first put forward the “spleen theory” of HAPC (Figure 7). It hypothesizes that HH exposure induced spleen ferroptosis, especially RPMs, which further inhibited erythrophagocytosis, RBC clearance, HGB processing and iron recycling. This theory identifies 637 a novel pathophysiological process in which HH exposure leads to the induction of spleen ferroptosis, particularly in red pulp. As a result, erythrophagocytosis, RBCs clearance, and iron recycling are impeded, causing a buildup of RBCs within the spleen, leading to splenomegaly and continued RBC production. The “spleen theory” of HAPC provides an intriguing new perspective on the physiological responses to hypoxia and highlights the role of the spleen as a vital organ in managing erythrocyte homeostasis under high altitude/hypoxic conditions. These findings may be clinically relevant to the pathological conditions of HAPC progression.

Revised version: Based on the above experiments, we propose the “spleen theory” of HAPC as depicted in Figure 10. It hypothesizes that HH exposure precipitates ferroptosis within the spleen, particularly affecting RPMs. This process, in turn, impedes erythrophagocytosis, RBCs clearance, HGB processing, and iron recycling. Consequently, this leads to an accumulation of RBCs within the spleen, exacerbating splenomegaly and perpetuating the production of RBCs under HH conditions. The “spleen theory” of HAPC provides a novel perspective on the physiological responses to hypoxia and highlights the role of the spleen as a vital organ in managing erythrocyte homeostasis under high altitude/hypoxic conditions. These findings may be clinically relevant to the pathological conditions of HAPC progression.

Added in the revised version: HA/HH exposure suppresses erythrophagocytosis by inducing macrophages ferroptosis in the spleen.

Added in the revised version: This figure summarizes the multifaceted impact of HA and HH exposure on splenic macrophage function. HA/HH exposure is demonstrated to trigger significant iron mobilization, notably increasing Fe2+ levels, which are further elevated in the spleen and macrophages of mice. Concurrently, HA/HH exposure is associated with an upregulation of ACSL4 expression, coupled with a suppression of xCT expression and GSH production. This alteration results in enhanced lipid peroxidation, culminating in the escalation of macrophage ferroptosis. Additionally, the exposure adversely affects monocyte migration and differentiation from BM to the spleen. Collectively, these effects of HA/HH exposure contribute to a notable inhibition of erythrophagocytosis, thereby promoting erythrocytosis and potentially exacerbating the progression of HAPC. This comprehensive analysis underscores the intricate interplay between environmental factors and cellular mechanisms in the spleen, which could significantly influence the pathogenesis of HAPC.

Abbreviations

Added in the revised version: bone marrow

Added in the revised version: Cysteine

Added in the revised version: Glutathione

Added in the revised version: Ferritin light chains

Added in the revised version: Ferroportin

Added in the revised version: Ferritin

Added in the revised version: Hypobaric hypoxia

Added in the revised version: High-altitude

Added in the revised version: High-altitude polycythemia

Added in the revised version: Heme oxygenase-1

Added in the revised version: Haematocrit

Added in the revised version: Haemoglobin

Added in the revised version: Hypoxia inducible factor 1 a

Added in the revised version: Mean corpuscular haemoglobin

Added in the revised version: Malondialdehyde

Added in the revised version: Normobaric normoxia

Added in the revised version: nuclear receptor coactivator 4

Added in the revised version: peripheral blood mononuclear cells

Added in the revised version: Red blood cell

Added in the revised version: Red pulp macrophage

Added in the revised version: Transferrin receptor 1

Added in the revised version: Transferrin

Added in the revised version: Transferrin receptor

Removed in the revised version: BM: bone marrow; Cys: Cysteine; GSH: Glutathione; Ft-L: Ferritin light chains; Fpn: Ferroportin; Ft: Ferritin; HH: Hypobaric hypoxia; HA: High-altitude; HAPC: High-altitude polycythemia; HO-1: Heme oxygenase-1; HCT: Haematocrit; HGB: Haemoglobin; HIF-1a: Hypoxia inducible factor 1 a; MCH: Mean corpuscular haemoglobin; MDA: Malondialdehyde; NN: Normobaric normoxia; NCOA4: nuclear receptor coactivator 4; PBMCs: peripheral blood mononuclear cells; RBC: Red blood cell; RPM: Red pulp macrophage; TfR1: Transferrin receptor 1; Tf: Transferrin; TfR: Transferrin receptor.

Author contributions

QQL and GHW conceived, organized, and designed the study. YQG supervised the work. WPY, MQL, JD and JYL performed the experiments. WPY, GW and BL contributed to the analysis of data. QQL and GHW prepared, wrote, and revised the manuscript. All authors contributed to read, and approved the submitted version.

Funding

Earlier version: This work was supported by Natural Science Foundation of China (Grants 32271228, 82171190, and 81873924), Open Cooperation Program from Key Laboratory of Extreme Environmental Medicine, Ministry of Education (KL2019GY011), China Postdoctoral Science Foundation (2020M673649), Cultivate Candidate of the Jiangsu Province “333” Project, and Nantong Municipal Science and Technology Project (MS12021020 and MS22021010).

Revised version: This work was supported by Natural Science Foundation of China (Grants 32271228, 81873924 and 82171190), Open Cooperation Program from Key Laboratory of Extreme Environmental Medicine, Ministry of Education (KL2019GY011), China Postdoctoral Science Foundation (2020M673649), Cultivate Candidate of the Jiangsu Province “333” Project.

Availability of data and materials

Earlier version: The datasets used and/or analyzed during the current study are available from the corresponding author on reasonable request.

Revised version: The datasets used and/or analysed during the current study are available from the corresponding author on reasonable request.

Declarations

Ethics approval and consent to participate

All animal care and experimental protocols were carried out according to the Chinese Animal Management Rules of the Ministry of Health and were authorized by the Animal Ethics Committees of Nantong University research program protocol #S20190219-011. The results of the study are presented clearly, honestly, and without fabrication, falsification, or inappropriate data manipulation.

Consent for publication

Not applicable.

Competing interests

The authors have no relevant financial or non-financial interests to disclose.

Author details

1Department of Physiology and Hypoxic Biomedicine, Institute of Special Environmental Medicine and Co-innovation Center of Neuroregeneration, Nantong University, 9 Seyuan Road, Chongchuan District, Nantong, Jiangsu 226019, China. 2College of High-Altitude Military Medicine, Institute of Medicine and Hygienic Equipment for High Altitude Region, Army Medical University, Chongqing 400038, China. 3Key Laboratory of Extreme Environmental Medicine and High-Altitude Medicine, Ministry of Education of China, Chongqing 400038, China. 4Department of Neurosurgery, Southwest Hospital, Army Medical University, Chongqing, Chongqing 400038, China.

Removed in the revised version: Supporting information